# 🏫 HW2: Exploratory Data Analysis

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

You must submit this assignment to Gradescope by the on-time deadline, **Monday, July 6th, 11:59 PM**. Please read the syllabus for the Slip Day policy. No late submissions beyond the details in the Slip Day policy will be accepted. While course staff is happy to help you if you encounter difficulties with submission, we may not be able to respond to late-night requests for assistance (TAs need to sleep, after all!). **We strongly encourage you to plan to submit your work to Gradescope several hours before the stated deadline.** This way, you will have ample time to contact staff for submission support. 

After completing this part ("Homework 2 Coding"), please submit the generated zip file to the Homework 2 Coding assignment on Gradescope. Gradescope will automatically submit a PDF of your written responses to the HW 2 Coding Written assignment; there is no need to submit it manually.

## 💪 Collaboration Policy

Data science is a collaborative activity. While you may talk with others about
the homework, we ask that you **write your solutions individually**. If you 
discuss the assignments with others, please **include their names** below.

**Collaborator(s)**: *List all collaborator(s) here*.

## 💯 Score Breakdown
Question | Manual? | Points
--- | --- | --
1 | Yes | 5
2 | Yes | 3
3a | Yes | 1
3b | Yes | 1
3c | Yes | 1
3d | Yes | 1
3e | Yes | 1
3f | Yes | 1
4a | No | 1
4b | No | 1
4c | No | 1
4d | Yes | 1
4e | Yes | 2
4f | No | 3
Total |   | 23

**Note**: "Manual" questions are written response questions that will be graded manually by the grading team instead of being graded by the autograder.

## ✊ Before You Start

### Autograder and Answer Cells

For each question in the assignment, please write down your answer in the answer cell(s) right below the question. 

We understand that it is helpful to have extra cells breaking down the process toward reaching your final answer. If you happen to create new cells *below* your answer to run code, **NEVER** add cells between a question cell and the answer cell below it. It may cause errors when we run the autograder, and it will cause a failure to generate the PDF file.



### Initialize your environment

The below cell should run without error if you're using the course DataHub.

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
```

<!-- BEGIN QUESTION -->

<br/><br/>

---
## 🗣️ Question 1: Plot presentation

In their 2013 research paper titled [The Missing "One-Offs": The Hidden Supply of High-Achieving, Low-Income Students](https://www.brookings.edu/wp-content/uploads/2016/07/2013a_hoxby.pdf), Caroline Hoxby and Chris Avery investigate the behavior of high-achieving low-income applicants to undergraduate programs in the United States.

Here is Figure 10 of their research paper:

<img src="assets/hoxby-average-fig10.png" width="500" alt="Stacked histogram showing that high-achieving low-income students apply more often to selective colleges where their scores are far above the college median, while high-income students apply more often to colleges closer to their score level.">

**Your task**: Record a 60-90 second screencast describing the plot above.
-   If you need more than 60-90 seconds to record for accessibility reasons, please let the teaching staff know via Ed.

Guidelines:

-   Your audience is a UC Berkeley undergraduate who has not taken a statistics or data science course, and is unfamiliar with the research paper. You do not need to read the research paper in full to understand the plot, but you might find it helpful to reference.
-   In your screencast, you must explain the minimum necessary background required to understand the plot, along with the key takeaway(s) of the plot.
-   Your name and face should not be in the screencast. In other words, the plot should take up the entire window of the screencast, with your voice playing in the background. **Do not introduce yourself or use any identifying information.**
-   You should use your mouse pointer to indicate particular points of interest on the plot. Alternatively, you can verbally direct the viewer to points of interest (e.g., "In the top right corner, you can see that...").
-   You should use an engaging tone that sounds as though you are presenting to a live audience.
-   You are welcome to read off a script that you have written ahead of time, but try not to make it too obvious. In fact, it can be very helpful to write a script for a presentation ahead of time, even if you do not actually read the script when presenting.

**Remember the three key guiding questions you should address before digging into the details: (1) What's the meaning of the X axis (e.g., what does a -10 value mean)? (2) What's the meaning of the Y axis? (3) What does a specific point/line/feature on your plot mean in context?**

There are lots of free tools for recording screencasts. For example, QuickTime is a useful tool for recording screencasts on a Mac. Feel free to post on Ed if you cannot identify a way to record a screencast anonymously.

As part of later homework, you will provide anonymous feedback on another student's submission for this problem.

-   Your screencast, along with the quality of feedback you provide to another student, will be graded by the teaching staff. **At no point will other students grade your work.**

💾 Upload your screencast to your UC Berkeley Google Drive account, or another place that allows you to share a public link.

-   You will submit a public link to your screencast in two places. See submission details at the top of this page.
-   Before submitting, make sure your screencast is publicly accessible with the link. One easy way to do this: Open the link in an incognito window in Google Chrome.
-   If your video is private or unviewable when we try to access it, we will not be able to confirm that you submitted your video before the HW2 deadline. You may lose points and/or have to use slip day(s).

**Why complete this problem?** Communication is a critical, but often under-appreciated, component of the data science life cycle.
This exercise helps develop your data storytelling ability, which is essential for getting anyone to actually care about your statistical analyses!

Include a public link to your screencast here. **Make sure to test your public link in an incognito window before submission.** Broken links will lose credit.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br/><br/>

---
## 🤖 Question 2: How well do LLMs do as analysts with minimal instructions?

Ask Gemini **Thinking Mode** or another **thinking** LLM model to answer the question we posed in class: 

> Scenario: You are an analyst working for the UC Berkeley admissions office.
> Your manager asks you to identify five California public high schools that defy typical patterns of application, admission, and/or attendance.
> The admissions office wants to visit these schools to figure out what's going really well or really wrong."

If you'd like, feel free to push the LLM to go in different directions or provide it with additional data or tips. 

What do you think about the answer provided by the LLM? Here are some questions to think about:
- What are the strengths and weaknesses of the analysis provided by the LLM, relative to the analysis we covered in lecture? 
- Did the LLM do a good job of "showing its work"? 
- Did the LLM defend its choices against alternatives using quantitative and/or qualitative reasoning?

Feel free to write your answer in sentences or bullet points. Your response should be the equivalent of at least three sentences.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 📝 Question 3: Making a complex plot

<img src="assets/asian-american-geography.png" width="1000" alt="Bubble scatter plot showing a slight negative relationship between states’ Asian-identifying share of high scorers and estimated attendance rates for white high scorers without legacy status.">

Consider the dataset used to produce the plot above. 

- The plot is Figure 3 of [this research paper](https://5harad.com/papers/college-admissions.pdf). 
- The caption will be helpful to read! 
- You might search through the paper to find definitions of words like "Ivy-11", "legacy", and "high-scorer".

**Note: We recommend that you do not use an LLM to help you with this problem.**

<!-- BEGIN QUESTION -->

### Question 3a

What is the minimum number of rows needed to reproduce the plot above? Explain your choice.

If you need to make any assumptions, state them.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 3b

What is the minimum number of columns needed to reproduce the plot? 

Additionally, provide a brief description of the data contained in each column.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 3c

What is the granularity of the dataset used to produce the plot above?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 3d

In your opinion, what is the granularity of the raw data that was "rolled-up" in order to produce the dataset plotted above? Try to get as fine-grained as possible. 

If you need to make any assumptions, state them.

*Hint: When is data about college applications first entered into an electronic database? You should use this first-stage data to answer the question.*

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 3e

What are the essential columns in the raw data from part (d) that are required to produce the summarized data plotted above?

If you need to make any assumptions, state them.

*Hint: When application data is first entered into an electronic database, what information is recorded? What subset of information is referenced in the plot above?*

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 3f

Write an analysis plan for transforming the raw data from parts (d) and (e) into the dataset used to make the plot above. Your analysis plan should contain at least one "For each..." statement. 

If you need to make any assumptions, state them.

*Hint: You will find it helpful to keep track of multiple intermediate datasets and join them. Break this problem into smaller pieces. For example, how do you calculate the values on the x-axis? How do you calculate the values on the y-axis? What about the other dimensions of the plot?*

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 🏫 Question 4

In this question, you will partially reproduce the cleaned college admissions dataset we used during lecture. 

- Note that we don't provide precise steps for every operation you'll carry out in this problem. To become a data scientist, you need to learn to be scrappy and successful with only high-level instructions. 
- Of course, if you're struggling to proceed, come to office hours or post on Ed. We're here to help!

*Note: You will find it helpful to use an LLM to assist you on this problem as needed. Remember, UC Berkeley students get free Gemini Pro access.*

```python

```

### Question 4a

 Head to the [UC Admissions by source school](https://www.universityofcalifornia.edu/about-us/information-center/admissions-source-school) page. Download the “Counts of fall freshmen by race/ethnicity” CSV data for UC Berkeley 2025 admits from California public high schools.

- Throughout this question, you are going to have to look up how to do things like `How do I read a CSV file using Pandas?` or `How do I count the number of rows in a dataframe?`. Scrappiness and self-help are key skills of a data scientist!

Assign the number of rows of the dataframe to the variable `q4a`.

```python
q4a = ...
```

```python
grader.check("q4a")
```

### Question 4b

Using an LLM to help you, or by searching the Pandas documentation, write Pandas code to remove all rows and columns from the Q4a dataset that contain data we did not use as part of our exploration in lecture. 

- Then, pivot the filtered data so it looks the same as the data on Slide 16 in Lecture 2. 

Finally, assign the number of rows of the resulting pivoted dataframe to the variable `q4b`.

- You can compare your data to a data preview from the slides to confirm that you have carried out this step correctly.

```python
q4b = ...
```

```python
grader.check("q4b")
```

### Question 4c

Next, head to the [California Department of Education website](https://www.cde.ca.gov/ds/) and download the data that contains the Grade 12 enrollment of each school in California. 

- Make sure to use the data from most appropriate academic year given that our UC admissions data is from 2025.

- Read the data documentation to figure out which enrollment counts you should use for each school. You will have a lot of options to choose from.

Before uploading the dataset to DataHub, **you will have to convert your downloaded data from `latin1` encoding to `utf-8` encoding.** 

- There is more than one way to do this. 

- One option is to upload the dataset into Google Sheets, and then download the resulting data in Google Sheets as a CSV. If you go with this method, you'll need to wait up to a minute for the data to finish uploading.

Finally, upload your [UTF-8](https://en.wikipedia.org/wiki/UTF-8) encoded dataset into DataHub. 

- Make sure that your uploaded data can be viewed as a table using DataHub's data viewer. You can open the data viewer by double clicking a `.csv` or `.tsv` file.

Assign the number of rows of this final dataframe to the variable `q4c`.

```python
...

```

```python
grader.check("q4c")
```

<!-- BEGIN QUESTION -->

### Question 4d

In the next few parts, we will use the provided `data/hs-names-translator.csv` file to help us join the admissions data to the California HS enrollment data. 

- If you're interested in learning how this translator file was produced, come to instructor office hours!

If we're trying to reproduce the analysis from lecture, does it make sense to use a left join, right join, full join, or inner join to merge the UC admission dataset and the California HS enrollment dataset? Defend your choice in at least one sentence.

```python
translator = pd.read_csv("data/hs-name-translator.csv")
translator.head(10)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 4e

In the lecture example, we had four schools with applications rates over 100%. One of them was "GIRLS ACADEMIC LEADERSHIP ACAD". Find this school in the translator file. 

Does it look like this school was matched correctly? If not, do you have a hypothesis as to why it could have been matched incorrectly? Write an explanation at least one sentence long.

```python

```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 4f

First, merge the UC admissions dataframe with the HS translator data using the appropriate type of join. 

- Then, merge the resulting dataframe with the California HS enrollment data using the appropriate type of join. 

- Finally, create a new column containing the application rate for each high school. 

To confirm that your merge was done correctly, you can reference any relevant data snippet shown in the lecture slides. 

- **The `grade_12` column used in the lecture dataset will have slightly different values than the correct answer to this question. As long as the values are close, you are all set!**

Assign the number of rows of your final dataframe to the variable `q4f`

```python
q4f = ...
```

```python
grader.check("q4f")
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 2 Coding!

### Course Content Feedback

If you have any feedback about this assignment or about any of our other weekly, weekly assignments, lectures, or discussions, please fill out the [Course Content Feedback Form](https://docs.google.com/forms/d/e/1FAIpQLSeNCHbTUipouL3yWFeRdl9Blwu-U4zSciTi4Kcb1p5ykCtb-A/viewform?usp=dialog). Your input is valuable in helping us improve the quality and relevance of our content to better meet your needs and expectations!

### Submission Instructions

Below, you will see one cell. Running this will automatically generate a zip file with your answers. Please submit this file to the Homework 2 Coding assignment on Gradescope. Gradescope will automatically submit the PDF from this file to the Homework 2 Coding Written assignment. **There is no need to manually submit Homework 2 Written answers; however, please check that the PDF was generated and submitted correctly**.

You are responsible for submitting the zip file generated below as well as a PDF with your answers to Homework 2 Math Prerequisites.
* **Homework 2 Coding**: Submit your Jupyter Notebook zip file for Homework 2 Coding,
which can be generated and downloaded from DataHub by using the `grader.export()`
cell provided below.
* **Homework 2 Coding Written**: Gradescope will automatically submit the PDF from the zip file submitted earlier. You do not need to submit anything to this assignment yourself, but please check that the submission went through properly.  

**You are responsible for ensuring your submission follows our requirements and that the automatic submission for Homework 2 Coding Written answers went through properly. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q4a (points: 1)
```
>>> assert q4a == 4671
```


Tests for q4b (points: 1)
```
>>> assert q4b == 1268
```


Tests for q4c (points: 1)
```
>>> assert q4c == 269158
```


Tests for q4f (points: 3)
```
>>> assert q4f == 1232
```
