# 🏘️ Project A2: Predicting Housing Prices in Cook County

## Due Date: Tuesday, July 28th 11:59 PM PDT
You must submit this assignment to Gradescope by the on-time deadline, Tuesday, July 28th 11:59 PM PDT. Please read the syllabus for the Slip Day policy. No late submissions beyond what is outlined in the Slip Day policy will be accepted. **We strongly encourage you to submit your work to Gradescope several hours before the stated deadline.** This way, you will have ample time to reach out to staff for support if you encounter difficulties with submission. While course staff is happy to help guide you with submitting your assignment ahead of the deadline, we will not respond to last-minute requests for assistance.

Please read the instructions carefully when you are submitting your work to Gradescope.


## Collaboration Policy

Data science is a collaborative activity. While you may talk with others about the project, we ask that you **write your solutions individually**. If you discuss the assignments with others, please **include their names** below.

**Collaborators:** *list names here*

```python
# Initialize Otter
import otter
grader = otter.Notebook("projA2.ipynb")
```

## [IMPORTANT] Coding Material for Data 100 Exams

As described in [Josh's Ed post](https://edstem.org/us/courses/99874/discussion/8172355), new coding material has been added to the scope of Data 100 exams. We will use some of this material in Project A2 to help you practice.

After reading the announcement and watching the recorded video, assign the passcode from the video to `passcode` below.

```python
passcode = ...
```

## Introduction

In Project A1, you performed some basic Exploratory Data Analysis (EDA), laying out the thought process that leads to certain modeling decisions. Then, you added a few new features to the dataset and cleaned the data.

In this project, we will specify and fit a linear model to a few features of the housing data to predict home prices. Then, we will analyze the error of the model and brainstorm ways to improve the model's performance. Finally, we'll think about the implications of predictive modeling within the Cook County Assessor's Office (CCAO) case study. This is especially relevant because statistical modeling is how the CCAO values properties. 

Given the history of racial discrimination in housing policy and property taxation in Cook County, consider the impacts of your modeling results as you work through this project, and think about what fairness might mean to property owners in Cook County.

After this part of the project, you should be comfortable with:
- Implementing a data processing pipeline using `pandas`.
- Using `scikit-learn` to build and fit linear models.

## Score Breakdown

Question | Manual | Points
----|----|----
Announcement| No | 1
1a | Yes | 1
1b | Yes | 1
1c | No | 1
1d | Yes | 1
1e | Yes | 1
2 | Yes | 2
3a | Yes | 2
3b | No | 3
3c | No | 2
4a | Yes | 2
4b | No | 1
4ci | Yes | 1
4cii | No | 1
5a | No | 0
5b | No | 0
5c | No | 0
5d | No | 3
5e | No | 0
5f | No | 0
6a | No | 1
6b | No | 2
6c | Yes | 2
7a | Yes | 1
7b | Yes | 2
8a | Yes | 2
8b | No | 0
Test Prediction | No | 3
Total | 12 | 36

## Before You Start

For each question in the assignment, please write down your answer in the answer cell(s) right below the question. 

We understand that it is helpful to have extra cells breaking down the process towards reaching your final answer. If you happen to create new cells below your answer to run code, **NEVER** add cells between a question cell and the answer cell below it. It will cause errors when we run the autograder, and it will sometimes cause a failure to generate the PDF file.

**Important note: The local autograder tests will not be comprehensive. You can pass the automated tests in your notebook but still fail tests in the autograder.** Please be sure to check your results carefully.

### Debugging Guide
If you run into any technical issues, we highly recommend checking out the [Data 100 Debugging Guide](https://ds100.org/debugging-guide/). In this guide, you can find general questions about Jupyter notebooks / Datahub, Gradescope, common `pandas` errors, `regex`, visualizations, and Proj. A1 and A2 common questions.

```python
import numpy as np

import pandas as pd
from pandas.api.types import CategoricalDtype

%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
from sklearn import linear_model as lm

import warnings
warnings.filterwarnings("ignore")

import zipfile
import os

from ds100_utils import run_linear_regression_test, run_linear_regression_test_optim
from feature_func import *

# Plot settings
plt.rcParams['figure.figsize'] = (12, 9)
plt.rcParams['font.size'] = 12
```

Let's load the training, validation, and test data.

```python
with zipfile.ZipFile('cook_county_data.zip') as item:
    item.extractall()
```

The dataset is split into a training set, a validation set, and a test set. In this project, you will train and validate a model using the training and validation sets, and then use a model trained on the combined training and validation sets to predict the **logarithm (log)** of `Sale Price`s from the test set.

Importantly, the test set we provide to you does **not** contain the values of our target variable, **log**(`Sale Price`). You can only use this test set to generate predictions.

In the cell below, we load the training and validation sets into the `DataFrame` `training_val_data` and the test set into the `DataFrame` `test_data`.

> If you are wondering why we are predicting **log(sale price)** instead of the sale price itself: In Project A1, we added the `Log Sale Price` column, which is the logarithm of the sale price column, to neutralize the impact of extreme right-skewed outliers on the fitted model. 
>
> However, even if your model fits well by accurately predicting the log sale price, that does **not** mean it will always perform well after exponentiating the prediction back to actual sale price. Here's a simple example to illustrate this:
>
> Let’s say the dataset of actual sale prices is { $e^0 = \$1$, $e^{10} \approx \$22,000$ }. In other words, we have one home worth $\$1$, and another home worth $\$22,000$.
>
> Suppose we train a constant model on log prices using MSE. The log prices are { $\log(e^0) = 0$, $\log(e^{10}) = 10$}.
>
> The model that would best predict the log prices {0, 10} will always output a log sale price of 5 (i.e., the average of 0 and 10). In other words, the model will always output a raw price of $e^5 \approx \$148$.
>
> Even though this model has minimized error on **log** prices, it has not minimized error on **actual** prices. If we predict $\$148$, we have a huge error on the home worth $\$22,000$!
>
> If we wanted to minimize error on actual prices, we should train the model on actual prices. If we were to train a constant model on actual prices, it would always predict $\frac{1+22000}{2} \approx 11000$.
>
> For the sake of accurately evaluating your model, we will be testing your predictions against the **log** of sale prices, rather than exponentiating them and comparing to the raw sale prices.
>
> That being said, we can still convert the predicted log sale price to the sale price to gain some intuition about the model and dataset, which we will explore in later questions.

```python
training_val_data = pd.read_csv("cook_county_train_val.csv", index_col='Unnamed: 0')
test_data = pd.read_csv("cook_county_contest_test.csv", index_col='Unnamed: 0')
from feature_func import remove_outliers

```

As a good sanity check, we should at least verify that the shape of the data matches the description.

```python
# 204792 observations and 62 features in training data
assert training_val_data.shape == (204792, 62)

# 55311 observations and 61 features in test data
assert test_data.shape == (55311, 61)

# Sale Price is provided in the training/validation data
assert 'Sale Price' in training_val_data.columns.values
training_val_data = remove_outliers(training_val_data, 'Sale Price', lower=499)
training_val_data["Log Sale Price"] = np.log(training_val_data["Sale Price"])

# Sale Price and Log Sale Price are hidden in the test data
assert 'Sale Price' not in test_data.columns.values
assert 'Log Sale Price' not in test_data.columns.values
```

Next, we print the columns in the dataset. 

> Remember, a more detailed description of each variable is included in `codebook.txt`, which is in the same directory as this notebook. 

```python
training_val_data.columns.values
```

<!-- BEGIN QUESTION -->

<br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 👥 Question 1: Human Context and Ethics

In this part of the project, we will explore the human context of our housing dataset. **You should watch [Lecture 14](https://ds100.org/su26/lectures/14/) before attempting this question.**

<br>

--- 

### Question 1a

Which of the following scenarios strikes you as unfair, and why? You can choose more than one. There is no single right answer, but you must explain your reasoning. Would you consider some of these scenarios more (or less) fair than others? Why?

A. A homeowner whose home is assessed at a higher price than it would sell for.  

B. A homeowner whose home is assessed at a lower price than it would sell for. 

C. An assessment process that systematically overvalues inexpensive properties and undervalues expensive properties.  

D. An assessment process that systematically undervalues inexpensive properties and overvalues expensive properties.

*Your response for each chosen scenario should be approximately 1 to 2 sentences.*

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

--- 

### Question 1b
Consider the following question: *"How much is a house worth?"*

In this question, we ask you to consider three specific parties that would be interested in the answer to this question.

Party 1: People selling homes.

Party 2: CCAO appraisers.

Party 3: People buying homes.

For each of the three parties above, state whether they have an interest in seeing a low or high housing price.

- **Hint**: Recall from Lecture 14 - The CCAO gets a fixed amount of money every year from property taxes, and ensures that home value assessments sum to this same total amount of money across all homes. This means that the tax assessment is zero-sum. Consider how this informs your answer for Party 2.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br>

---

### Question 1c

Consider a model fit to $n = 50$ training observations. Which of the following three residual plots corresponds to a model that makes property assessments resulting in **regressive** taxation? Recall from Lecture 14 that regressive taxation overvalues inexpensive properties and undervalues expensive properties. 

Assign `q1c` to the string letter corresponding to your plot choice.

We denote the response as $y$ (Log Sale Price), the prediction as $\hat{y}$, and the corresponding residual as $y - \hat{y}$. Assume that all three plots use the same vertical scale and that the horizontal line marks $y - \hat{y} = 0$. 

**Hint:** When a model overvalues a property (predicts a `Sale Price` greater than the actual `Sale Price`), how do $y$ and $\hat{y}$ compare? What about when a model undervalues a property?

<img src='images/res_plots.png' width="900px" alt="Each plot has y on the x-axis and y-y_hat on the y-axis. In plot A, as the x-values increase so do the y-values. In plot B, there is no relationship between the x-values and the y-values. In plot C, as the x-values increase, the y-values decrease." />

```python
q1c = ...
```

```python
grader.check("q1c")
```

## The CCAO Dataset

You'll work with the dataset from the Cook County Assessor's Office (CCAO) in Illinois. This government institution determines property taxes across most of Chicago's metropolitan areas and nearby suburbs. In the United States, all property owners must pay property taxes, which are then used to fund public services, including education, road maintenance, and sanitation. These property tax assessments are based on property values estimated using statistical models considering multiple factors, such as real estate value and construction cost.

However, this system is not without flaws. In late 2017, a lawsuit was filed against the office of Cook County Assessor Joseph Berrios for producing "[racially discriminatory assessments and taxes](https://www.chicagotribune.com/politics/ct-cook-county-board-assessor-berrios-met-20170718-story.html)." The lawsuit included claims that the assessor's office undervalued high-priced homes and overvalued low-priced homes, creating a visible divide along racial lines. Wealthy homeowners, who were typically white, paid less in property taxes, whereas [working-class, non-white homeowners paid more](https://www.chicagotribune.com/news/breaking/ct-cook-county-assessor-berrios-sued-met-20171214-story.html).

The Chicago Tribune's four-part series, "[The Tax Divide](https://www.chicagotribune.com/investigations/ct-tax-divide-investigation-20180425-storygallery.html)," delves into how this was uncovered. After "compiling and analyzing more than 100 million property tax records from the years 2003 through 2015, along with thousands of pages of documents, then vetting the findings with top experts in the field," they discovered that "residential assessments had been so far off the mark for so many years." [You can read more about their investigation here.](https://apps.chicagotribune.com/news/watchdog/cook-county-property-tax-divide/assessments.html)

Make sure to watch [Lecture 14](https://ds100.org/su26/lectures/14/) before answering the following questions!

<!-- BEGIN QUESTION -->

<br>

---

### Question 1d

Consider the parties from `1b` who have opposing perspectives. When their interests conflict, whose perspective should the CCAO prioritize and why? (There is no single right answer to this question, but in order to receive full credit, you must come up with a convincing answer that is grounded in what you learned in Lecture 14).

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

--- 

### Question 1e

What were the central problems with the earlier property tax system in Cook County as reported by the Chicago Tribune? What were the primary causes of these problems? 

*Your response should be approximately 2 to 4 sentences.*

**Note:** Along with reading the paragraph above, you will need to watch [Lecture 14](https://ds100.org/su26/lectures/14/) to answer this question.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

<hr style="border: 1px solid #fdb515;" />

## 🧑‍🍳 Question 2: Preparing Data


We will next split the dataset into a training set and a validation set. We will use the training set to fit our model's parameters and the validation set to evaluate how well our model will perform on unseen data drawn from the same distribution as the training set. 

> If we used all the data to fit our model, we would not have a way to estimate model performance on **unseen data** such as the test set in `cook_county_contest_test.csv`.

In the code cell below, the implementation of `train_val_split` has been provided for you. For each designated line, write a comment explaining what the line does and how it contributes to constructing the training or validation set.


[Take a look at the documentation](https://numpy.org/doc/stable/reference/random/generated/numpy.random.permutation.html) for `np.permutation` to better understand what `shuffled_indices` contains.

```python
# This makes the train-validation split in this section reproducible across different runs 
# of the notebook. You do not need this line to run train_val_split in general.

# DO NOT CHANGE THIS LINE
np.random.seed(1337)
# DO NOT CHANGE THIS LINE

def train_val_split(data):
    """ 
    Takes in a DataFrame `data` and randomly splits it into two smaller DataFrames 
    named `train` and `validation` with 80% and 20% of the data, respectively. 
    """
    
    data_len = data.shape[0]
    shuffled_indices = np.random.permutation(data_len)

   
    # 1. Write your explanation here.
    train_indices = shuffled_indices[:int(data_len * 0.8)]
    # 2. Write your explanation here.
    validation_indices = shuffled_indices[int(data_len * 0.8):]
    # 3. Write your explanation here.
    train = data.iloc[train_indices]
    # 4. Write your explanation here.
    validation = data.iloc[validation_indices]

    

    return train, validation
train, validation = train_val_split(training_val_data)
```

<!-- END QUESTION -->

<br>

<hr style="border: 1px solid #fdb515;" />

## 🎛️ Question 3: Fitting a Simple Model

Let's fit a linear regression model using the ordinary least squares (OLS) estimator. We will start with a simple model that uses only two features: the **number of bedrooms** in the household and the **log-transformed total area covered by the building** (in square feet). 

Consider the following expression for our first linear model that contains one of the features:

$$
\text{Log Sale Price} = \theta_0 + \theta_1 \cdot (\text{Bedrooms})
$$

In parallel, we will also consider a second model that contains both features:

$$
\text{Log Sale Price} = \theta_0 + \theta_1 \cdot (\text{Bedrooms}) + \theta_2 \cdot (\text{Log Building Square Feet})
$$

<!-- BEGIN QUESTION -->

<br>

--- 

### Question 3a

**Without running any calculation or code**, assign `q3a` to the comparator ('>=', '>', '=', '<', '<=') that correctly completes the following statement below.

Then, explain your choice in one or two sentences in the comment provided in the answer cell.

We quantify the loss on our linear models using MSE (Mean Squared Error). Consider the training loss of the first model and the training loss of the second model. If we use normal equations to obtain the optimal set of parameters, we are guaranteed that:

$$
\text{Training Loss of the 1st Model} ~~  \_\_\_\_ ~~ \text{Training Loss of the 2nd Model}
$$

```python
q3a = ...
# Write your explanation here.
```

```python
grader.check("q3a")
```

<!-- END QUESTION -->

Let's now see how we might create model 1 in code. Recall that model 1 makes the following prediction:

$$
\text{Log Sale Price} = \theta_0 + \theta_1 \cdot (\text{Bedrooms})
$$

As you may recall from Project A1, information such as the numbers of bathrooms and bedrooms is embedded in a text field rather than stored in separate columns.

In Project A1, you implemented an `add_total_bathrooms` function to extract the number of bathrooms. Here, we provide an analogous implementation, `add_total_bedrooms`, that extracts the number of bedrooms.

```python
from feature_func import add_total_bedrooms
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error

training_set_with_total_bedrooms = add_total_bedrooms(train)
X_train_m1 = training_set_with_total_bedrooms[["Bedrooms"]]

# Take a look at the input to the model used for predicting log sale price
display(X_train_m1.head())

# Fit the model
model_m1 = LinearRegression()
model_m1.fit(X_train_m1, train["Log Sale Price"])
```

Let's make one small tweak to our approach. Rather than generate the X variable for our model using two lines of code floating around on their own, let's generate the X variable using a function:

```
def model_one_pipeline(df):
    training_set_with_total_bedrooms = add_total_bedrooms(df)
    X_m1 = training_set_with_total_bedrooms[["Bedrooms"]]
    return X_m1
```

This will be useful for a couple of reasons:
1. It will give us a standard code pattern for creating models.
2. It will allow us to apply this function to the training (for fitting) and validation set (for evaluating).

For example, consider the code below which uses this function to fit a model then uses it again to compute the MSE on the validation set.

```python
def model_one_pipeline(df):
    # Create new DataFrame by applying add_total_bedrooms to df, our input DataFrame
    training_set_with_total_bedrooms = add_total_bedrooms(df)
    # Extract only feature data from our new DataFrame
    X_m1 = training_set_with_total_bedrooms[["Bedrooms"]]
    return X_m1

# Apply our newly created pipeline function to the train DataFrame to create our design matrix containing the “Bedroom” feature
X_train_m1 = model_one_pipeline(train)
# Create a LinearRegression object and assign it to model_m1
model_m1 = LinearRegression()
# Fit our new LinearRegression object to our design matrix
model_m1.fit(X_train_m1, train["Log Sale Price"])
# Apply our pipeline function to the original validation set in order to extract the “Bedroom” feature our model uses for prediction
X_validation_m1 = model_one_pipeline(validation)
# Use our transformed validation set to generate Log Sale Price predictions, and then find MSE for these predictions compared to actual Log Sale Price for the validation set
mean_squared_error(model_m1.predict(X_validation_m1), validation["Log Sale Price"])
```

**Note:** If you have prior experience with modeling libraries such as scikit-learn, you might wonder whether it is possible to use a scikit-learn pipeline. Yes! It would look something like this:

```
# Create the pipeline
pipeline = Pipeline([
    ('add_total_bedrooms', FunctionTransformer(add_total_bedrooms)),
    ('select_bedrooms', FunctionTransformer(lambda df: df[['Bedrooms']])),
    ('regression_model', LinearRegression()),    
])
pipeline.fit(train, train['Log Sale Price'])
```

However, this approach can get a bit messy to use later on this assignment, so we're going to stick with our own custom pipeline functions for the rest of this assignment. 

<br>

--- 

### Question 3b

Now that we've shown you how to build model 1, you'll now build model 2 and compare the performance of the two models to confirm your answer to part 3a. As a refresher, the equations for model 1 and model 2, respectively, are:
$$
\text{Log Sale Price} = \theta_0 + \theta_1 \cdot (\text{Bedrooms})
$$


$$
\text{Log Sale Price} = \theta_0 + \theta_1 \cdot (\text{Bedrooms}) + \theta_2 \cdot (\text{Log Building Square Feet})
$$


Assign the final training data and validation data for both models to the following set of variables:

- First Model: `X_train_m1`, `X_valid_m1`. This is already implemented for you. 
- Second Model: `X_train_m2`, `X_valid_m2`. Please implement this in the second cell below. You may use the first model as an example.

**Hint:** `np.log(...)` may be useful. The expression
`df[["a", "b"]]` returns a `DataFrame` containing columns `a` and `b`.
Use `.copy()` before adding a new column so that the selected data can
be modified independently.

```python
def model_one_pipeline(df):
    data_with_total_bedrooms = add_total_bedrooms(df)
    X_m1 = data_with_total_bedrooms[["Bedrooms"]]
    return X_m1
    
X_train_m1 = model_one_pipeline(train)
X_valid_m1 = model_one_pipeline(validation)
```

```python
def model_two_pipeline(df):
    data_with_total_bedrooms = ...
    X_m2 = data_with_total_bedrooms[...].copy()
    X_m2["Log Building Square Feet"] = ...
    return X_m2[...]


X_train_m2 = model_two_pipeline(train)
X_valid_m2 = model_two_pipeline(validation)

# Take a look at the result
display(X_train_m2.head())
display(X_valid_m2.head())
```

To test your code, run the cells below:

```python
grader.check("q3b")
```

<br>

--- 

### Question 3c

Finally, let's do some regression!

We first initialize a [`sklearn.linear_model.LinearRegression` object (documentation)](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LinearRegression.html) for both of our models. We set the `fit_intercept = True` so that the model estimates an intercept term rather than forcing the intercept to zero.

```python
linear_model_m1 = lm.LinearRegression(fit_intercept=True)
linear_model_m2 = lm.LinearRegression(fit_intercept=True)
```

Now it's time to fit our linear regression model. Use the cell below to fit both models. Then, for each fitted model, compute the predicted values of `Log Sale Price` for the training data and the predicted values of `Log Sale Price` for the validation data.

Assign the predicted values from both of your models on the training and validation set to the following variables:

- For the first Model, assign the predicted values for the **training set** to `m1_training_yhat`, and predicted values for the **validation set** to `m1_valid_yhat`.
- For the second model, assign the predicted values for the **training set** to `m2_training_yhat`, and the predicted values for the **validation set** to `m2_validation_yhat`.

**Note**: To make sure you understand how to compute the correct predicted values for both the training and validation data set, there won't be any hidden tests for this part.
                                                                            
**Hint:** The [`.predict(X)` method](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LinearRegression.html#sklearn.linear_model.LinearRegression.predict)
will be helpful here. You may also refer to this
[official example showing how a fitted model generates predictions](https://scikit-learn.org/1.9/auto_examples/linear_model/plot_ols_ridge.html#model-evaluation).

```python
# Fit the 1st model
...
# Compute the fitted and predicted values of Log Sale Price for 1st model
m1_training_yhat = ...
m1_valid_yhat = ...

# Fit the 2nd model
...
# Compute the fitted and predicted values of Log Sale Price for 2nd model
m2_training_yhat = ...
m2_validation_yhat = ...
```

```python
grader.check("q3c")
```

<br>

<hr style="border: 1px solid #fdb515;" />

## 📊 Question 4: Evaluate Our Simple Model

<br>

--- 

Let's now move into the analysis of our two models!

```python
def rmse(predicted, actual):
    """
    Calculates RMSE from actual and predicted values.
    Input:
      predicted (1D array): Vector of predicted/fitted values
      actual (1D array): Vector of actual values
    Output:
      A float, the RMSE value.
    """
    return np.sqrt(np.mean((actual - predicted)**2))
```

<!-- BEGIN QUESTION -->

<br>

--- 

### Question 4a

We can assess a model's performance and quality of fit with a plot of the residuals ($y-\hat{y}$) versus the observed outcomes ($y$). 

In the code cell below, begin with comments that provide an analysis plan for constructing a residual plot for Model 2 on the validation set. Your plan should clearly identify:
1. The values that should be plotted on the horizontal axis.
2. How the residuals on the vertical axis will be calculated.
3. How to add a horizontal reference line at a residual of zero.
4. The plot type, the axis labels, and the title.

Then, use an LLM to translate your analysis plan into code that produces the visualization. You do not need to create a separate DataFrame unless doing so makes your code clearer.
We will grade the resulting visualization. Your plot should have an informative title and labeled axes.

```python
# Write your analysis plan here.
```

<!-- END QUESTION -->

<br>

--- 

### Question 4b

Based on the structure you see in your plot, does this model appear to correspond to _regressive_, _fair_, or _progressive_ taxation?

Assign "regressive", "fair" or "progressive" to `q4b` in the cell below accordingly.


```python
q4b = ...
```

```python
grader.check("q4b")
```

<!-- BEGIN QUESTION -->

<br>

--- 

### Question 4c, Part i

`k`-fold cross validation (CV) allows us to more accurately evaluate our models than a single train-validation split. In `k`-fold CV, every data point is used in model training at some point. 

> Heads up: You should use `k`-fold cross validation to evaluate the different models you create in Question 5, as you will have limited submissions per day on Gradescope to assess your test RMSE. 

In this Project, we won't use `k`-fold CV to optimize for hyperparameters like we will in Lecture 15, since OLS does not have any hyperparameters.

Instead, we will use `k`-fold cross validation to generate `k` estimates of the out-of-sample RMSE. These `k` RMSEs might be helpful when deciding between models in Question 5.   

Below, we provide an implementation of `compute_CV_error`. Fill in
the designated comments to explain the key steps.

The function uses [`KFold.split`](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.KFold.html#sklearn.model_selection.KFold.split) to generate the training and validation
indices for each fold. For every split, it fits the model on the
training fold, computes RMSE on the validation fold, and appends that
RMSE to a list. Finally, it returns the list of validation RMSEs.
    
**Note:** The tests for this subpart only check the return value of `compute_CV_error` is a list of the correct length, 4cii will check for correctness.

```python
from sklearn.model_selection import KFold

def compute_CV_error(X_train, Y_train, folds=10):
    """
    Split the training data into `k` subsets.
    For each subset, 
        - Fit a model holding out that subset.
        - Compute the RMSE on that subset (the validation set).
    You should be fitting `k` models in total.
    Return a list of `k` RMSEs.

    Args:
        X_train (DataFrame): Training data.
        Y_train (Series): Label.
        folds (int): Number of cross-validation folds.
    
    Return:
         A list of `k` RMSEs.
    """
    model = lm.LinearRegression(fit_intercept=True) # Use model that Question 5 requires.
    kf = KFold(n_splits=folds)
    validation_errors = []
    
    for train_idx, valid_idx in kf.split(X_train):



        # 1. Write your explanation here.
        split_X_train, split_X_valid = X_train.iloc[train_idx], X_train.iloc[valid_idx]
        split_Y_train, split_Y_valid = Y_train.iloc[train_idx], Y_train.iloc[valid_idx]
      
      
        # 2. Write your explanation here.
        model.fit(split_X_train, split_Y_train)

        # 3. Write your explanation here.
        error = rmse(model.predict(split_X_valid), split_Y_valid)


        validation_errors.append(error)

    return validation_errors
    
```

```python
grader.check("q4ci")
```

<!-- END QUESTION -->

<br>

--- 

### Question 4c, Part ii

Use `compute_CV_error` to generate two lists of RMSEs generated by 4-fold cross validation: one list for Model 1 and one list for Model 2. Assign these lists to `cv_m1` and `cv_m2`, respectively. 

To ensure the tests run as expected and the correct answers align with the tests, please do not change the random seed used in this question.

```python
# DO NOT CHANGE THIS LINE
np.random.seed(1337)

# MODIFY THESE LINES
cv_m1 = ...
cv_m2 = ...
print(f"The RMSE errors for 4-fold cross-validation on Model 1 were: {cv_m1}")
print(f"The RMSE errors for 4-fold cross-validation on Model 2 were: {cv_m2}")
```

```python
grader.check("q4cii")
```

While our simple model explains some of the variability in price, there is certainly still a lot of room for improvement —— one reason is we have been only using 1 or 2 features (out of a total of 70+)! In the next part, you will engineer and incorporate more features to improve the model's fairness and accuracy.

<br>

<hr style="border: 1px solid #fdb515;" />

## 🏗️ Question 5

It is time to build your own model! Here's what's next:

1. You will conduct feature engineering on your training data using the `feature_engine_final` function (defined in `q5d`), and you will fit and evaluate the model with the training data.
2. You will process the test data with `feature_engine_final`, use your trained model to predict `Log Sale Price` for the test set, and save those predictions in the required CSV format for submission.

Your goal in Question 5 is to:

* Define a function to perform feature engineering and produce a design matrix for modeling.
* Apply this feature engineering function to the training data and use it to train a model that can predict the `Log Sale Price` of houses.
* Use this trained model to predict the `Log Sale Price` values for the test set. Remember that the test set does not contain the true `Sale Price` of any house—your model can only generate predictions!
* Submit your predicted `Log Sale Price` values for the test set to Gradescope.


Right under the grading scheme, we will outline some important Datahub logistics. **Please make sure you read this carefully to avoid running into memory issues later!**

* Question 5a (Not graded): You will explore possible features for your model.
* Question 5b (Not graded): You will perform EDA on the dataset. 
* Question 5c (Not graded): You can define feature engineering helper functions. 
* Question 5d (**Graded!**): You will create your design matrix and train a model. 
* Question 5e (Not graded): You will fit and evaluate your model. 
* Question 5f (**Graded!**): You will generate the predictions for the test set.

### Question 5 Grading Scheme

Your grade for Question 5 will be based on your model's RMSE when making predictions on the training set, as well as your model’s RMSE when making predictions on the test set. The tables below provide scoring guidelines. If your RMSE lies in a particular range, you will receive the number of points associated with that range.

**Important**: While your training RMSE can be checked at any time in this notebook, your test RMSE can only be checked by submitting your model’s predictions to Gradescope. **You will only be able to submit your test set predictions to Gradescope up to 4 times per day**. Attempts will not carry over across days, so we recommend planning ahead to make sure you have enough time to fine-tune your model! 

The thresholds are as follows:

Points | 3 | 2 | 1 | 0
--- | --- | --- | --- | ---
Training RMSE | Less than 0.65 | [0.65, 1.0) | [1.0, 1.75) | More than 1.75

Points | 3 | 2 | 1 | 0
--- | --- | --- | --- | ---
Test RMSE | Less than 0.75 | [0.75, 1.2) | [1.2, 1.9) | More than 1.9

An RMSE of 1 on the natural log scale is equivalent to a typical **multiplicative error** of $e \approx 2.718$. 

In other words, an RMSE of 1 on the log scale could be *roughly* interpreted as follows: Typically, our predictions are different from actual prices by a factor of 2.718. 

Example intuition: Let's say for one data point our actual home price is $e^{12} \approx \$162,500$. The predicted home price is $e^{13} \approx 442,000$. On the log scale, the actual home price is 12, and the predicted price is 13. So, our log error is 1. This is equivalent to a multiplicative error of $\frac{\$442,000}{\$162,500} \approx e \approx 2.718$. 

<br><br>

---
### Some notes before you start

- **If you are running into memory issues, restart the kernel and only run the cells you need to.** The cell below (question cell) contains most to all of the imports necessary to successfully complete this portion of the project. You do not need to rerun the cell at the top of this notebook to complete Question 5. The autograder will have more than 4GB of memory, so you will not lose credit as long as your solution to Question 5 is within the total memory (4GB) limits of Datahub. By default, we reset the memory and clear all variables using `%reset -f`. If you want to delete specific variables, you may also use `del` in place of `%reset -f`. For example, the following code will free up memory from data used for older models: `del training_val_data, test_data, train, validation, X_train_m1, X_valid_m1, X_train_m2, X_valid_m2`. Our staff solution can be run independently from all other questions, so we encourage you to do the same to make debugging easier.
- **If you need the data again after deleting the variables or resetting, you must reload them again from earlier in the notebook.**
- You will be predicting `Log Sale Price` on the data stored in `cook_county_contest_test.csv`. 
- You MUST remove any **new cells** you add before submitting to Gradescope to avoid any autograder errors. 
- **You can only submit your test set prediction CSV file to Gradescope up to 4 times per day. Start early!** In the case that you are approved for an extension, you will be granted 4 more submissions for each day the deadline has been extended.

**PLEASE READ THE ABOVE MESSAGE CAREFULLY!**

```python
# The 3 lines below to clean up memory from previous questions and reinitialize Otter!
# If you want to refer to any functions or variables you defined at any point earlier in the project, 
# Place them in the cell under Question 5c so that you can access them after the memory is reset.
# If you think you will not run into any memory issues, you are free to comment out the next 3 lines as well.

%reset -f                               
import otter                            
grader = otter.Notebook("projA2.ipynb")
```

```python
# Imports all the necessary libraries again

import numpy as np
import pandas as pd
from pandas.api.types import CategoricalDtype

%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
from sklearn import linear_model as lm

import warnings
warnings.filterwarnings("ignore")

import zipfile
import os

from ds100_utils import run_linear_regression_test, run_linear_regression_test_optim
from feature_func import *

from sklearn.preprocessing import OneHotEncoder
```

<br>

---

### Question 5a: Finding Potential Features

**This question is not graded** – it is intended to give helpful guidance on how to get started with feature engineering in `q5d`. You may write as little or as much as you would like here; it will not factor into your grade. 

Read the documentation about the dataset in `codebook.txt`, located in this directory. Is there any data you think may be related to housing prices? You may want to take note of them below for future reference.

<br>

---

### Question 5b: More EDA


**This question is not graded** – it is intended to give helpful guidance on how to get started with feature engineering. You may write as little or as much as you would like here; it will not factor into your grade. 

Use the scratch space below to conduct any additional EDA you would like to see. You may use this space to make additional plots to help you visualize the relationship between variables or compute any relevant summary statistics. You are free to add any number of cells as needed below and before the next question. You may find it helpful to review Project A1 and the techniques we explore there.

Some potential ideas:
    
* Plot the distribution of a variable. Is this variable heavily skewed? Are there any outliers? This can inform how you engineer your features later on.

* Make a scatter plot between quantitative features and the outcomes. Is there a relationship? Is there a transformation that may linearize the relationship?

* Make a plot of a qualitative feature and the outcome. Is there a relationship? How can we transform this categorical data into numerical features that can be useful for OLS?

* Find the correlation coefficient between features and the outcome. Is there a strong relationship between the two? Can you find the correlation coefficient between different transformations of the feature and the outcome?


Choose one or more visualizations that will help you decide which features or transformations may be useful in your final model.
Before writing code, describe each intended visualization in comments. For each visualization, your analysis plan should identify:
1. The input column or columns.
2. Any filtering, imputation, or transformation you will apply.
3. The relationship or distribution the visualization will display.
4. How the result may inform feature engineering.
Then, use an LLM to translate your analysis plan into code. This question is not graded, but your work should be reproducible and clearly labeled.

```python
# Write your analysis plan here.
# Add any EDA code below
```

<br>

---

### Question 5c: Defining Helper Function or Helper Variables

**This question is not graded, but we suggest that you put all your helper functions below for readability and ease of testing.** Use this space below to define any additional helper functions you may use in your final model. These can be transformation functions you identified in the optional question above. 

We also provide several helper functions for you to use. You may also use the functions provided earlier in this question and in Project A1.

**Important:** Your use of LLMs for parts (c) and (d) must be limited to the uses permitted by our [course syllabus](https://ds100.org/su26/syllabus/#collaboration-policy-and-academic-honesty). You **may not** use an LLM to write code for you. Please come to office hours if you need help.

```python
# Define any additional helper functions or variables you need here

def ohe(data, labels):
    """
    One-hot encode the specified columns.
    """
    assert isinstance(labels, list), "labels must be a list of columns"

    encoder = OneHotEncoder()
    encoder.fit(data[labels])

    dummies = pd.DataFrame(
        encoder.transform(data[labels]).todense(),
        columns=encoder.get_feature_names_out(),
        index=data.index
    )

    return data.join(dummies), encoder.get_feature_names_out()

def bin_col(data, colname, granularity):
    """
    Put a numerical column into categorical bins.
    """
    data = data.copy()

    min_value = data[colname].min()

    data[colname + " Bin"] = (
        (data[colname] - min_value) // granularity
    ).astype(str)

    return data

```

<br>

---

## Question 5d: Defining The Pipeline Function

Just as in the guided model from `q3b`, you'll be writing a function that generates features for your model. Your function should be called `final_model_pipeline` in the following cell. It should be just like `model_one_pipeline` and `model_two_pipeline` in that it takes as input a training or validation set, and outputs the appropriate columns to be provided to a model.

Note: The `final_model_pipeline` function starts with special code that ignores house prices above a fixed threshold. We've chosen a threshold of 1,000,000 dollars. Note that we already discarded all of the one dollar houses at the very beginning of the notebook, since this was just a placeholder value, so we don't need to remove these again. The reason we're setting this upper limit is that there are a small number of very expensive properties. Since the MSE is the squared difference between prediction and actual value, this means these properties will have a relatively large impact on the loss, and thus our model is going to try too hard to make these properties' predictions accurate. Since the model pays so much attention to only the most expensive houses, it will do worse on the much more common lower priced household.

You're welcome to experiment with a different threshold. 

To get a really rough sense of the quality of your final_model_pipeline, run the cell immediately below, which will print the RMSE on the training set. You should be able to get your RMSE on the training set below 0.65 on the Log Sale Price scale.

```python
from feature_func import remove_outliers

# Please include all of your feature engineering processes inside this function.
# Do not modify the parameters of this function.
def final_model_pipeline(df): 
    # make a copy to avoid modifying the original dataframe
    data = df.copy()

    # We can see the sale price so this must be a training set. Test sets will not have
    # the Sale Price column available. Let's avoid fitting house prices that are either
    # very low or very high. Do not remove this code, as keeping all the one dollar houses
    # and super expensive houses will affect your model negatively.
    if "Sale Price" in data.columns:
        data = remove_outliers(data, 
                               variable='Sale Price',
                               upper=1_000_000)

        # drop the sale price if present to avoid accidentally using these in the model
        data = data.drop(["Sale Price", "Log Sale Price"], axis='columns')
    

```

```python
# Run this cell before you get to next step.
def make_feature_engine_adapter(pipeline):
    """
    Adapt final_model_pipeline(df) -> X to the A2 interface:

        feature_engine_final(train_data) -> (X, y)
        feature_engine_final(test_data, is_test_set=True) -> X
    """
    training_columns = None

    def feature_engine_final(data, is_test_set=False):
        nonlocal training_columns

        working_data = data.copy()

        if is_test_set:

            pipeline_input = working_data.drop(
                columns=["Sale Price", "Log Sale Price"],
                errors="ignore"
            )

            X = pipeline(pipeline_input)


            if len(X) != len(working_data):
                raise ValueError(
                    "The pipeline removed rows from validation/test data."
                )

            if not X.index.equals(working_data.index):
                raise ValueError(
                    "The pipeline changed or reordered validation/test indices."
                )

            if training_columns is None:
                raise RuntimeError(
                    "Call feature_engine_final on training data before "
                    "processing validation/test data."
                )


            X = X.reindex(columns=training_columns, fill_value=0)

            return X


        if "Sale Price" not in working_data.columns:
            raise ValueError(
                "Training data must contain a Sale Price column."
            )


        working_data = working_data.loc[
            working_data["Sale Price"] > 499
        ].copy()


        working_data["Log Sale Price"] = np.log(
            working_data["Sale Price"]
        )


        y = working_data["Log Sale Price"].copy()
        X = pipeline(working_data)


        y = y.loc[X.index]

        if not X.index.equals(y.index):
            raise ValueError("X and y are not aligned.")

        training_columns = X.columns.copy()

        return X, y

    return feature_engine_final


feature_engine_final = make_feature_engine_adapter(
    final_model_pipeline
)
```

```python
check_rmse_threshold = run_linear_regression_test_optim(
    lm.LinearRegression(fit_intercept=True),
    feature_engine_final,
    "cook_county_train.csv",
    None,
    False
)

print(
    "Current training log RMSE:",
    check_rmse_threshold.loss
)
```

```python
grader.check("q5d")
```

<br>

---

### Question 5e: Fit and Evaluate your Model

**This question is not graded.** Use this space below to evaluate your models. Some ideas are listed below. 

**Note:** For grading purposes, you should use a `sklearn` default `LinearRegression()` model with intercept term. Do not modify any hyperparameter in `LinearRegression()`, and focus on feature selection or hyperparameters of your own feature engineering function.

It may also be helpful to calculate the RMSE directly as follows:

$$RMSE = \sqrt{\dfrac{\sum_{\text{houses in the set}}(\text{actual log price for house} - \text{predicted log price for house})^2}{\text{number of houses}}}$$

A function that computes the RMSE is provided below. Feel free to use it if you would like to calculate the RMSE for your training set.

```python
def rmse(predicted, actual):
    """
    Calculates RMSE from actual and predicted values.
    Input:
      predicted (1D array): Vector of predicted/fitted values
      actual (1D array): Vector of actual values
    Output:
      A float, the RMSE value.
    """
    return np.sqrt(np.mean((actual - predicted)**2))
```

<details>
    <summary>[<b>Click to Expand</b>] Hints: </summary>
    
Train set:

* Check your RMSE. Is this a reasonable number? You may use our grading scheme as a reference. Keep in mind that training error is generally less than testing error. Recall that an error of 1 on a natural log scale is equivalent to a **multiplicative error** of $e \approx 2.718$. In other words, an RMSE of 1 on the log scale could be *roughly* interpreted as follows: Typically, our predictions are different from actual prices by a factor of 2.718. For example, if we predict $\$1000$, the actual home price might be around $\$1000*2.718=\$2718$ or $\$1000/2.718=\$368$.

Test set:
* Find the original data shape at the beginning of the notebook (in the provided assert statement). What should the output shape be?

* Since test and training/validation sets come from the same population (recall that test and training/validation sets are a random split from larger data), we expect our test prediction to have a similar range as the validation data. Plot the observed training (Log) Sale Price and the predicted (Log) Sale Price. Are the ranges similar? Do you have any unreasonable extreme prediction that cannot be exponentiated?

* We cannot compute test RMSE directly since we do not have the observed values. Perform cross-validation (feel free to use `compute_CV_error` from 4c) to estimate your test error. Recall that we are treating the validation set as unseen data.
    
</details>

```python
# Use this space to evaluate your model
# if you reset your memory, you need to define the functions again
```

<br>

---

### Question 5f: Submission

**This question is graded.**

Recall that the test set provided to you in `cook_county_contest_test.csv` does not contain values for the true `Sale Price` or `Log Sale Price` of each house. You can only compute predictions of `Log Sale Price`. 

To determine your model's RMSE on the test set, you will submit the predictions made by your model to Gradescope. There, we will run checks to see what your test RMSE is by considering the (hidden) true values of `Log Sale Price` for the observations in the test set. 

Your score on this section will be determined by the grading scheme outlined at the start of Question 5. **Remember that you can only submit your test set predictions to Gradescope up to 4 times per day. Plan your time to ensure that you can adjust your model as necessary, and please test your model's performance using cross-validation before making any submissions.** 

- For more on cross-validation, check [Lecture 15](https://ds100.org/su26/lectures/15/). Feel free to use `compute_CV_error` from 4c as well! 

To determine the error on the test set, please submit your predictions on the test set to the Gradescope assignment **Project A2 Test Set Predictions**. The CSV file to submit is generated below, and you should not modify the cell below. Simply download the CSV file, and submit it to the appropriate Gradescope assignment.

**You will not receive credit for the test set predictions (i.e., up to 3 points) unless you submit to this assignment**!!

**Note:** If you run into any errors, the [Proj. A2 Common Mistakes](https://ds100.org/debugging-guide/proja2/) section of the [Data 100 Debugging Guide](https://ds100.org/debugging-guide) may be a helpful resource.

```python
from datetime import datetime
from IPython.display import display, HTML

Y_test_pred = run_linear_regression_test(lm.LinearRegression(fit_intercept=True), feature_engine_final, None, 'cook_county_train.csv', 'cook_county_contest_test.csv', 
                                         is_test = True, is_ranking = False, return_predictions = True
                                         )

# Construct and save the submission:
submission_df = pd.DataFrame({
    "Id": pd.read_csv('cook_county_contest_test.csv')['Unnamed: 0'], 
    "Value": Y_test_pred,
}, columns=['Id', 'Value'])
timestamp = datetime.now().strftime("%Y%m%d_%H%M%S")
filename = "submission_{}.csv".format(timestamp)
submission_df.to_csv(filename, index=False)

#print('Created a CSV file: {}.'.format("submission_{}.csv".format(timestamp)))
display(HTML("Download your test prediction <a href='" + filename + "' download>here</a>."))
print('You may now upload this CSV file to Gradescope for scoring.')#
```

```python
# Scratch space to check if your prediction is reasonable. See 5e for hints. 
# We will not reset the submission count for mis-submission issues.
submission_df["Value"].describe()
```

Congratulations on finishing your prediction model for home sale prices in Cook County! In the following section, we'll delve deeper into the implications of predictive modeling within the CCAO case study, especially because statistical modeling is how the CCAO valuates properties. 

Refer to [Lecture 14](https://ds100.org/su26/lectures/14/) if you're having trouble getting started!

<br>

<hr style="border: 1px solid #fdb515;" />

## 🫚 Question 6: Exploring RMSE

In this question, we will explore what RMSE means in the context of predicting house prices. We will go through different ways of visualizing the performance of your model and see how it relates to questions about property taxes. 

In the cell below, we create the `preds_df` `DataFrame` for use in this question.

<br>

---

```python
# Run the cell below; no further action is needed
train_df = pd.read_csv('cook_county_train.csv')
X, Y_true = feature_engine_final(train_df)
model = lm.LinearRegression(fit_intercept=True)
model.fit(X, Y_true)
Y_pred = model.predict(X)
```

```python
preds_df = pd.DataFrame({'True Log Sale Price' : Y_true, 'Predicted Log Sale Price' : Y_pred, 
                         'True Sale Price' : np.e**Y_true, 'Predicted Sale Price' : np.e**Y_pred})
preds_df.head()
```

<br>

---
### Question 6a

Let's examine how our model performs on two halves of our data: `cheap_df`, which contains the rows of `preds_df` with prices below or equal to the median sale price, and `expensive_df`, which has rows of `preds_df` with true sale prices above the median. 

Take a moment to understand what is happening in the cell below, as it will also prove useful in `q6b`.  

```python
# Run the cell below to obtain the two subsets of data; no further action is needed.
min_Y_true, max_Y_true = np.round(np.min(Y_true), 1) , np.round(np.max(Y_true), 1)
median_Y_true = np.round(np.median(Y_true), 1)
cheap_df = preds_df[(preds_df['True Log Sale Price'] >= min_Y_true) & (preds_df['True Log Sale Price'] <= median_Y_true)]
expensive_df = preds_df[(preds_df['True Log Sale Price'] > median_Y_true) & (preds_df['True Log Sale Price'] <= max_Y_true)]

print(f'\nThe lower interval contains houses with true sale price ${np.round(np.e**min_Y_true)} to ${np.round(np.e**median_Y_true)}')
print(f'The higher interval contains houses with true sale price ${np.round(np.e**median_Y_true)} to ${np.round(np.e**max_Y_true)}\n')
```

**Compute the RMSE of your model's predictions of `Sale Price` (i.e., the exponentiated version of `Log Sale Price`) on `cheap_df` and `expensive_df` separately**. Assign these RMSE values to `rmse_cheap` and `rmse_expensive`, respectively.

Then, compute the proportion of houses in `cheap_df` and `expensive_df` for which the model overestimates `Sale Price`. Assign the proportions to `prop_overest_cheap` and `prop_overest_expensive`, respectively. 

- For example, if we were working with a dataset of 3 houses where the actual `Log Sale Price`s were [10, 11, 12] and the model predictions were [5, 15, 13], then the proportion of houses with overestimated values would be 2/3.

**Note:** When calculating `prop_overest_cheap` and `prop_overest_expensive`, you could use either `Log Sale Price` or `Sale Price`. Take a second to think through why this metric is unchanged under a log transformation.

**Note:** As mentioned before, even though it is not mathematically correct to compare the RMSE in sale price—since we trained our model using log sale price—it still provides some insight into the dataset and our model.

```python

rmse_cheap = rmse(..., ...)
rmse_expensive = rmse(..., ...)

prop_overest_cheap = np.sum(...) / len(...)
prop_overest_expensive = np.sum(...) / len(...)
print(f"The RMSE for properties with log sale prices in the interval {(min_Y_true, median_Y_true)} is {np.round(rmse_cheap)}")
print(f"The RMSE for properties with log sale prices in the interval {(median_Y_true, max_Y_true)} is {np.round(rmse_expensive)}\n")
print(f"The percentage of overestimated values for properties with log sale prices in the interval {(min_Y_true, median_Y_true)} is {np.round(100 * prop_overest_cheap, 2)}%")
print(f"The percentage of overestimated values for properties with log sale prices in the interval {(median_Y_true, max_Y_true)} is {np.round(100 * prop_overest_expensive, 2)}%")
```

```python
grader.check("q6a")
```

<br>

---
### Question 6b

The intervals we defined above are broad. Let's try and take a more fine-grained approach to understand how RMSE of sale price and proportion of houses overestimated vary across different intervals of `Log Sale Price`. 

The functions `rmse_interval` and `prop_overest_interval` compute the RMSE of **sale price (NOT log sale price)** and proportion of overestimated home prices for any interval of home prices, respectively. In the cell below, complete the implementation of these two functions.

Pay close attention to the function descriptions, and feel free to reuse and modify the code you wrote in the previous part as needed.

**Note:** The autograder tests provided for each of the functions are **not** comprehensive as the outputs of the function are highly dependent on your model. Make sure that the values you obtain are interpretable and that the plots in 6c look appropriate.

```python

def rmse_interval(df, start, end):
    """
    Given a DataFrame df, computes the RMSE of Sale Price for rows whose
    True Log Sale Price lies in the interval [start, end].

    Input:
    df : pandas DataFrame with columns 'True Log Sale Price',
        'Predicted Log Sale Price', 'True Sale Price', 'Predicted Sale Price'
    start : A float specifying the start of the interval (inclusive)
    end : A float specifying the end of the interval (inclusive)
    """

    subset_df = df[...]
    rmse_subset = rmse(..., ...)
    return rmse_subset

def prop_overest_interval(df, start, end):
    """
    Given a DataFrame df, computes the proportion of overestimated houses
    whose True Log Sale Price lies in the interval [start, end].

    Input:
    df : pandas DataFrame with columns 'True Log Sale Price',
        'Predicted Log Sale Price', 'True Sale Price', 'Predicted Sale Price'
    start : A float specifying the start of the interval (inclusive)
    end : A float specifying the end of the interval (inclusive)
    """

    subset_df = df[...]

    # DO NOT MODIFY THESE TWO LINES
    if subset_df.shape[0] == 0:
        return -1

    prop_subset = np.sum(...) / len(...)
    return prop_subset
```

```python
grader.check("q6b")
```

<!-- BEGIN QUESTION -->

<br>

---
### Question 6c

Using the functions above, we can generate visualizations of how the RMSE of sale price and proportion of overestimated houses vary for different intervals:

```python
# RMSE plot
plt.figure(figsize = (8,5))
plt.subplot(1, 2, 1) 
rmses = []
for i in np.arange(8, 14, 0.5):
    rmses.append(rmse_interval(preds_df, i, i + 0.5))
plt.bar(x = np.arange(8.25, 14.25, 0.5), height = rmses, edgecolor = 'black', width = 0.5)
plt.title('RMSE of Sale Price For Different Intervals\n of Log Sale Price', fontsize = 10)
plt.xlabel('Log Sale Price')
plt.yticks(fontsize = 10)
plt.xticks(fontsize = 10)
plt.ylabel('RMSE')

# Overestimation plot  
plt.subplot(1, 2, 2)
props = []
for i in np.arange(8, 14, 0.5):
    props.append(prop_overest_interval(preds_df, i, i + 0.5) * 100) 
plt.bar(x = np.arange(8.25, 14.25, 0.5), height = props, edgecolor = 'black', width = 0.5)
plt.title('Percentage of House Values Overestimated \n for different intervals of Log Sale Price', fontsize = 10)
plt.xlabel('Log Sale Price')
plt.yticks(fontsize = 10)
plt.xticks(fontsize = 10)
plt.ylabel('Percentage of House Values\n that were Overestimated (%)')

plt.tight_layout()
plt.show()
```

Which of the two plots above would be more useful in ascertaining whether the assessments tended to result in progressive or regressive taxation? Provide a brief explanation to support your choice of plot.

Then, explain whether your chosen plot aligns more closely with scenario C or scenario D from `q1a`: 

    C. An assessment process that systematically overvalues inexpensive properties and undervalues expensive properties.  
    D. An assessment process that systematically undervalues inexpensive properties and overvalues expensive properties.

*Your response should be approximately 3 to 4 sentences.*

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

<hr style="border: 1px solid #fdb515;" />

## 🕵️ Question 7: Evaluating the Model in Context

<br>

---
### Question 7a

When evaluating your model, we used RMSE. In the context of estimating the value of houses, what does the residual mean for an individual homeowner? How does a positive or negative residual affect them in terms of property taxes? Discuss the cases where the residual is positive and negative separately.

*Your response should be approximately 2 to 4 sentences.*

_Type your answer here, replacing this text._

<!-- END QUESTION -->

In the case of the Cook County Assessor’s Office, Chief Data Officer Rob Ross states that fair property tax rates are contingent on whether property values are assessed accurately —— that they’re valued at what they’re worth, relative to properties with similar characteristics. This implies that having a more accurate model results in fairer assessments. The goal of the property assessment process for the CCAO, then, is to be as accurate as possible. 

When the use of algorithms and statistical modeling has real-world consequences, we often refer to the idea of fairness as a measurement of how socially responsible our work is. Fairness is incredibly multifaceted: Is a fair model one that minimizes loss - one that generates accurate results? Is it one that utilizes "unbiased" data? Or is fairness a broader goal that takes historical contexts into account?

These approaches to fairness are not mutually exclusive. If we look beyond error functions and technical measures of accuracy, we'd not only consider _individual_ cases of fairness but also what fairness —— and justice —— means to marginalized communities on a broader scale. We'd ask: What does it mean when homes in predominantly Black and Hispanic communities in Cook County are consistently overvalued, resulting in proportionally higher property taxes? When the white neighborhoods in Cook County are consistently undervalued, resulting in proportionally lower property taxes? 

Having "accurate" predictions doesn't necessarily address larger historical trends and inequities, and fairness in property assessments in taxes works beyond the CCAO's valuation model. Disassociating accurate predictions from a fair system is vital to approaching justice at multiple levels. Take Evanston, IL —— a suburb in Cook County —— as an example of housing equity beyond just improving a property valuation model: their City Council members [recently approved reparations for African American residents](https://www.usnews.com/news/health-news/articles/2021-03-23/chicago-suburb-approves-government-reparations-for-black-residents).


<!-- BEGIN QUESTION -->

<br>

---

### Question 7b

Reflecting back on your exploration in Questions 6 and 7a, in your own words, what makes a model's predictions of property values for tax assessment purposes "fair"? 

This question is open-ended and part of your answer may depend on your specific model; we are looking for thoughtfulness and engagement with the material, not correctness. 

**Hint:** Some guiding questions to reflect on as you answer the question above: What is the relationship between RMSE, accuracy, and fairness as you have defined it? Is a model with a low RMSE necessarily accurate? Is a model with a low RMSE necessarily "fair"? Is there any difference between your answers to the previous two questions? And if so, why?

*Your response should be approximately 1 to 2 paragraphs. Feel free to answer the questions in the hint to structure your answer.*

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

<hr style="border: 1px solid #fdb515;" />

## 🕵️ Question 8: Finding Better Metrics For Our Model

<br>

---

### Question 8a

As discussed in Project A1, RMSE—while a widely used and powerful error metric introduced in earlier lectures—is not always the most appropriate choice when evaluating the *fairness* of a property appraisal system. In Question 7, you already encountered some of RMSE's limitations, particularly its tendency to disproportionately emphasize errors in high-priced properties due to the squaring of residuals.

In this question, rather than relying on RMSE, we will train and evaluate our model using another custom fairness metric. Specifically, we will examine the **Mean Absolute Percentage Error (MAPE)**, which measures the average error as a percentage of the true value. This allows us to better assess whether the model is making relatively fair predictions across different segments of the housing market.

In the code cells that follow, we’ll explore how MAPE varies across price ranges—comparing the model’s relative performance on inexpensive versus expensive housing. This helps identify whether the model systematically favors or disadvantages certain groups of properties based on their price.

**Note:** You'll notice that we’re no longer using `lm.LinearRegression()`, but instead have started using `minimize`. This is because `scikit-learn` does not allow customization of the loss function in their standard linear regression model. The approach shown below provides an equivalent way to train a linear regression model using a custom error metric.

**Warning:** These cells take quite a long time to run. Please be patient, wait, and avoid restarting the kernel or rerunning these cells more than necessary.

```python
data = pd.read_csv("cook_county_train.csv", index_col='Unnamed: 0')
trainX, trainY = feature_engine_final(data)
# DO NOT MODIFY ANYTHING IN THIS CELL
np.random.seed(1337)
rows = np.random.choice(range(trainX.shape[0]), 3000, replace = False)
trainX = trainX.iloc[rows]
trainY = trainY.iloc[rows]
```

```python
# X is the design matrix (including bias column), y is the vector of true outputs, theta is the parameter vector
def mape(theta, X, y):
    y_pred = X @ theta  # compute predicted values using linear combination of features
    percentage_error = np.abs((y - y_pred) / y)  # calculate element-wise percentage errors
    return np.mean(percentage_error)  # return the MAPE
```

```python
from scipy.optimize import minimize

# Add bias (intercept) column to the design matrix
trainX_with_bias = np.column_stack([np.ones(trainX.shape[0]), trainX])

# Initialize parameter vector with zeros
theta_0 = np.zeros(trainX_with_bias.shape[1])

# Use scipy's minimize to find weights that minimize MAPE
res = minimize(mape, theta_0, args=(trainX_with_bias, trainY), method='Nelder-Mead')

# Optimal weights after training
theta_opt = res.x

theta_opt
```

```python
new_preds_df = pd.DataFrame({
    'True Log Sale Price'     : trainY,
    'Predicted Log Sale Price': trainX_with_bias @ theta_opt,
    'True Sale Price'         : np.e ** trainY,
    'Predicted Sale Price'    : np.e ** (trainX_with_bias @ theta_opt)
})

plt.figure(figsize=(8, 5))
plt.subplot(1, 2, 1)

mape_values = []
for i in np.arange(8, 14, 0.5):
    mape_values.append(mape_interval(new_preds_df, i, i + 0.5))

plt.bar(x=np.arange(8.25, 14.25, 0.5), height=mape_values, edgecolor='black', width=0.5)
plt.title('MAPE of Log Sale Price Across\nLog Sale Price Intervals', fontsize=10)
plt.xlabel('Log Sale Price')
plt.ylabel('Mean Absolute Percentage Error (MAPE)')
plt.xticks(fontsize=10)
plt.yticks(fontsize=10);
```

What can you infer from this graph using our new custom error metric?  Write your findings in the space below in three to five sentences. Consider addressing the following points:

- How does this metric differ from RMSE?  
- What is the purpose of this custom error function?  
- How does it relate to the idea of underpriced expensive housing mentioned in Question 6?
- Why would this potentially be better than the RMSE in terms of the CCAO dataset?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

### Question 8b (Bonus!!!)

**Note:** This question is not worth any points—it’s purely for those who are curious about applying their custom error metric to the CCAO dataset.

After seeing how the course staff implemented MAPE, you might want to test your own custom error metric and see how it performs on the dataset! We’ve implemented most of the code for you (similar to how MAPE was implemented above), and your only task is to complete the `your_custom_error_metric` and `your_custom_error_metric_interval` functions.  

To understand what’s expected, refer to the `mape` function in Question 8a and the `rmse_interval` function in Question 6b.

```python
def your_custom_error_metric(theta, X, y): # feel free to include more parameters
    ...
```

```python
def your_custom_error_metric_interval(df, start, end):
    ...
```

```python
from scipy.optimize import minimize

trainX_with_bias = np.column_stack([np.ones(trainX.shape[0]), trainX])
theta_0 = np.zeros(trainX_with_bias.shape[1])
# if you adjusted the amount of parameters above, you would have to modify the args parameter as well
res = minimize(your_custom_error_metric, theta_0, args=(trainX_with_bias, trainY), method='BFGS')
theta_opt = res.x
```

```python
new_preds_df = pd.DataFrame({
    'True Log Sale Price'     : trainY,
    'Predicted Log Sale Price': trainX_with_bias @ theta_opt,
    'True Sale Price'         : np.e ** trainY,
    'Predicted Sale Price'    : np.e ** (trainX_with_bias @ theta_opt)
})

plt.figure(figsize=(8, 5))
plt.subplot(1, 2, 1)

your_values = []
for i in np.arange(8, 14, 0.5):
    your_values.append(your_custom_error_metric_interval(new_preds_df, i, i + 0.5))

plt.bar(x=np.arange(8.25, 14.25, 0.5), height=your_values, edgecolor='black', width=0.5)
plt.title('Custom Error Metric of Sale Price Across\n Log Sale Price Intervals', fontsize=10)
plt.xlabel('Log Sale Price')
plt.ylabel('Custom Error Metric')
plt.xticks(fontsize=10)
plt.yticks(fontsize=10);
```

As a final thought, how did your custom error metric perform? In what sense, if any, is it “fairer” than the RMSE metric?

<br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## Course Content Feedback

If you have any feedback about this assignment or about any of our other weekly assignments, lectures, or discussions, please fill out the [Course Content Feedback Form](https://docs.google.com/forms/d/e/1FAIpQLSeNCHbTUipouL3yWFeRdl9Blwu-U4zSciTi4Kcb1p5ykCtb-A/viewform). Your input is valuable in helping us improve the quality and relevance of our content to better meet your needs and expectations!

## Submission Instructions

Below, you will see a cell. Running this cell will automatically generate a zip file with your autograded answers. Once you submit this file to the Project A2 Coding assignment on Gradescope, Gradescope will automatically submit a PDF file with your written answers to the Project A2 Written assignment. If you run into any issues when running this cell, feel free to check this [section](https://ds100.org/debugging-guide/autograder-gradescope/#why-does-the-last-grader-export-cell-fail-if-all-previous-tests-passed) in the Data 100 Debugging Guide.

If there are issues with automatically generating the PDF, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit those. 

**Please make sure you submit the following to the right assignments:**

* **Project A2 Coding:** Submit the zip file generated by using the `grader.export()` cell provided below.
* **Project A2 Written:** Gradescope will automatically submit the PDF from the zip file submitted earlier. You do not need to submit anything to this assignment yourself, but *please check that the submission went through properly and that all plots rendered correctly*.
* **Project A2 Test Set Predictions:** Submit the CSV file generated in `q5f`.

**You are responsible for ensuring your submission follows our requirements and that everything was generated and submitted correctly. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline. 

```python
# To ensure autograder functions correctly, run this cell before you run the last cell.
%reset -f 
import otter                            
grader = otter.Notebook("projA2.ipynb")

# Imports all the necessary libraries again

import numpy as np
import pandas as pd
from pandas.api.types import CategoricalDtype

%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
from sklearn import linear_model as lm

import warnings
warnings.filterwarnings("ignore")

import zipfile
import os

from ds100_utils import run_linear_regression_test, run_linear_regression_test_optim
from feature_func import *
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for announcement (points: 1)


Tests for q1c (points: 1)
```
>>> q1c.lower() in ['a', 'b', 'c']
True
```


Tests for q3a (points: 2)
```
>>> (q3a in ['>=', '=', '<=', '>', '<']) == True
True
```
```
>>> bool(q3a == '>=')
True
```


Tests for q3b (points: 3)
```
>>> assert isinstance(X_train_m1, pd.core.frame.DataFrame) and isinstance(X_valid_m1, pd.core.frame.DataFrame) and isinstance(X_train_m2, pd.core.frame.DataFrame) and isinstance(X_valid_m2, pd.core.frame.DataFrame)
```
```
>>> assert len(X_train_m2) == 135144
>>> assert len(X_valid_m2) == 33787
```
```
>>> assert np.all(X_train_m2['Bedrooms'].iloc[0:3] == [1, 3, 3])
>>> assert np.all(np.isclose(X_train_m2['Log Building Square Feet'].iloc[0:3], [6.216606, 7.340836, 7.445418]))
```


Tests for q3c (points: 2)
```
>>> assert np.isclose(m1_training_yhat.max(), 17.512560351964424, atol=0.0001) == True
```
```
>>> assert np.isclose(m2_training_yhat.max(), 15.603149674252519, atol=0.0001) == True
```
```
>>> assert np.isclose(m1_valid_yhat.max(), 15.53165404144502, atol=0.0001) == True
```
```
>>> assert np.isclose(m2_validation_yhat.max(), 14.852612886805643, atol=0.0001) == True
```


Tests for q4b (points: 1)
```
>>> bool(q4b.lower() in ['regressive', 'fair', 'progressive'])
True
```


Tests for q4ci (points: 1)
```
>>> len(compute_CV_error(X_train_m1, train['Log Sale Price'], folds=4)) == 4
True
```
```
>>> len(compute_CV_error(X_train_m1, train['Log Sale Price'], folds=6)) == 6
True
```


Tests for q4cii (points: 1)
```
>>> np.random.seed(1337)
>>> assert np.isclose(sum(cv_m1), 3.616638933868188, atol=0.0001)
```
```
>>> np.random.seed(1337)
>>> assert np.isclose(sum(cv_m2), 3.2268529515389055, atol=0.0001)
```


Tests for q5d (points: 3)
```
>>> bool(check_rmse_threshold(1.75))
True
```
```
>>> bool(check_rmse_threshold(1))
True
```
```
>>> bool(check_rmse_threshold(0.65))
True
```


Tests for q6a (points: 1)
```
>>> bool(rmse_cheap >= 0)
True
```
```
>>> bool(rmse_expensive >= 0)
True
```
```
>>> bool(0 <= prop_overest_cheap <= 1)
True
```
```
>>> bool(0 <= prop_overest_expensive <= 1)
True
```


Tests for q6b (points: 2)
```
>>> bool(prop_overest_interval(preds_df, 10, 14) >= 0 and prop_overest_interval(preds_df, 10, 14) <= 1)
True
```
```
>>> bool(rmse_interval(preds_df, 10, 14) < 1000000000.0 and rmse_interval(preds_df, 10, 14) > 0)
True
```
