# 🏡 Project A1: Exploring Cook County Housing

## 📅 Due Date: Tuesday, July 21 11:59 PM PDT
You must submit this assignment to Gradescope by the on-time deadline, Tuesday, July 21 11:59 PM PT. Please read the syllabus for the Slip Day policy. No late submissions beyond what is outlined in the Slip Day policy will be accepted. **We strongly encourage you to submit your work to Gradescope several hours before the stated deadline.** This way, you will have ample time to reach out to staff for support if you encounter difficulties with submission. While course staff is happy to help guide you with submitting your assignment ahead of the deadline, we will not respond to last-minute requests for assistance.

Please read the instructions carefully when you are submitting your work to Gradescope.

## Collaboration Policy

Data science is a collaborative activity. While you may talk with others about the project, we ask that you **write your solutions individually**. If you discuss the assignments with others, please **include their names** below.

**Collaborators:** *list names here*

```python
# Initialize Otter
import otter
grader = otter.Notebook("projA1.ipynb")
```

<hr style="border: 1px solid #fdb515;" />
<hr style="border: 5px solid #003262;" />
<br/><br/>

## Introduction

This project explores what can be learned from an extensive housing dataset that is embedded in a dense social context in Cook County, Illinois.

In Project A1 (this assignment), we will guide you through some basic Exploratory Data Analysis (EDA) to understand the structure of the data. Next, you will be adding a few new features to the dataset, while cleaning the data as well in the process. Lastly, you will design your own custom error metric to replace RMSE in this setting.

In Project A2 (the following assignment), you will specify and fit a linear model for the purpose of prediction. Finally, we will analyze the model's error, brainstorm ways to improve its performance, and test the custom error metric you defined in Project A1.


### 💯 Grading
Grading is broken down into auto-graded answers and free responses. For auto-graded answers, the results of your code are compared to provided and/or hidden tests. For free response, readers will evaluate how well you answered the question and/or fulfilled the requirements of the question. 

#### 🎼 Score breakdown

Question | Manual | Points
----|----|----
1a | Yes | 1
1bi | Yes | 1
1bii | Yes | 1
1ci | Yes | 1
1cii | Yes | 1
1d | Yes | 1
2a | No | 1
2b | Yes | 1
3a | Yes | 1
3b | Yes | 1
4a | Yes | 1
4b | Yes | 1
5a | No | 1
5b | No | 2
5c | Yes | 2
6a | No | 2
6b | No | 2
6c | No | 1
7a | Yes | 1.5
7b | Yes | 1.5
Total | 12 | 25

## 🫵 Before You Start

For each question in the assignment, please write down your answer in the answer cell(s) right below the question. 

We understand that it is helpful to have extra cells breaking down the process towards reaching your final answer. If you happen to create new cells below your answer to run code, **NEVER** add cells between a question cell and the answer cell below it. It will cause errors when we run the autograder, and it will sometimes cause a failure to generate the PDF file.

**Important note: The local autograder tests will not be comprehensive. You can pass the automated tests in your notebook but still fail tests in the autograder.** Please be sure to check your results carefully.

### 🐛 Debugging Guide
If you run into any technical issues, we highly recommend checking out the [Data 100 Debugging Guide](https://ds100.org/debugging-guide/). In this guide, you can find general questions about Jupyter notebooks / Datahub, Pensive, common `pandas` errors, `regex`, visualizations, and more.

```python
import numpy as np

import pandas as pd

%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns

import warnings
warnings.filterwarnings("ignore")

import zipfile
import os

# Plot settings
plt.rcParams['figure.figsize'] = (12, 9)
plt.rcParams['font.size'] = 12
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 📊 The Data

The dataset consists of over 500,000 records from Cook County, Illinois, the county where Chicago is located. The dataset has 61 features in total; the 62nd is `Sale Price`, which you will predict with linear regression in the next part of this project. An explanation of each variable can be found in the included `codebook.txt` file (you can optionally open this by first clicking the `data` folder, then clicking `codebook.txt` file in the navigation pane). Some of the columns have been filtered out to ensure this assignment doesn't become overly long when dealing with data cleaning and formatting.

The data are split into training and test sets with 204,792 and 68,264 observations, respectively, but we will only be working on the training set for this part of the project.

Let's first extract the data from the `cook_county_data.zip`. Notice we didn't leave the `csv` files directly in the directory because they take up too much space without some prior compression. Just run the cells below to load in the initial data: 

```python
with zipfile.ZipFile('data/cook_county_data.zip') as item:
    with item.open("cook_county_train.csv") as f:
        initial_data = pd.read_csv(f, index_col='Unnamed: 0')
```

Let's load the initial data.

As a good sanity check, we should at least verify that the data shape matches the description.

```python
# 204,792 observations and 62 features in training data
assert initial_data.shape == (204792, 62)
# Sale Price is provided in the training data
assert 'Sale Price' in initial_data.columns.values
```

The next order of business is getting a feel for the variables in our data.  A more detailed description of each variable is included in `codebook.txt` (in the same directory as this notebook).  **You should take some time to familiarize yourself with the codebook before moving forward.**

Let's take a quick look at all the current columns in our initial data.

```python
initial_data.columns.values
```

```python
initial_data['Description'][0]
```

<br/><br/>
<hr style="border: 1px solid #fdb515;" />

## 🗣️ Part 1: Contextualizing the Data

Let's look at the background of our dataset before diving into a full-scale analysis. 

<!-- BEGIN QUESTION -->

<br><br>

---

### 🚀 Question 1a

Based on the columns in this dataset and the values that they take, what do you think each row represents? That is, what is the **granularity** of this dataset? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

---
### 🚀 Question 1b

In this question, you will consider the types of questions you can answer with this dataset.

<!-- BEGIN QUESTION -->

#### Question 1bi

Craft at least two questions about housing in Cook County that can be answered with this dataset and provide the type of analytical tool you would use to answer it (e.g. "I would create a ___ plot of ___ and ___" or "I would calculate the ___ [summary statistic] for ___ and ____"). Be sure to reference the columns that you would use and any additional datasets you would need to answer that question.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### Question 1bii
For one of the questions that you asked in part `1bi`, draw by hand what a plot resulting from this analysis might look like. The trend does not have to be in line with what the actual data would show. Make sure your plot includes a title and axis labels.

Include a caption that answers your question from the previous part with the graph.

Submit this to the Project A1 Plot Sketch assignment on Gradescope, alongside your answer to `1cii`. Set `q1bii` to `True` once you've done so!

```python
q1bii = ...
```

```python
grader.check("q1bii")
```

<!-- BEGIN QUESTION -->

<br><br>

---
### 🚀 Question 1c

In this question, you will explore what types of questions you can answer with access to additional columns.

#### Question 1ci

Suppose now, in addition to the information already contained in the dataset, you also have access to several new columns containing **demographic data about the owner**, including race/ethnicity, gender, age, annual income, and occupation. Provide one new question about housing in Cook County that can be answered using at least one column of demographic data and at least one column of existing data and provide the type of analytical tool you would use to answer it.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### Question 1cii
For the question that you asked in part `1ci`, draw by hand what a plot resulting from this analysis might look like. You can assume the data look any way you would like, but you should be sure the hypothetical data is in realistic units. Make sure your plot includes a title and axis labels.

Include a caption that answers your question from the previous part with the graph.

Submit this to the Project A1 Plot Sketch assignment on Gradescope, alongside your answer to `1bii`. Set `q1cii` to `True` once you've done so!

```python
q1cii = ...
```

```python
grader.check("q1cii")
```

<!-- BEGIN QUESTION -->

<br><br>

---
### 🚀 Question 1d

Look at `codebook.txt` to see some of the unique regional features CCAO utilizes, such as `O’Hare Noise`. Now imagine you were in charge of predicting the `Sale Price` of houses in **your hometown** (your actual real life hometown/city - not the data provided). Propose a feature that you would want to collect specific to your location and hypothesize why it might be useful in predicting the sale price of houses.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/>
<hr style="border: 1px solid #fdb515;" />

## 🤔 Part 2: Exploratory Data Analysis

This dataset was collected by the [Cook County Assessor's Office](https://datacatalog.cookcountyil.gov/Property-Taxation/Archive-Cook-County-Assessor-s-Residential-Sales-D/5pge-nu6u) in order to build a model to predict the monetary value of a home. You can read more about data collection in the CCAO’s [Residential Data Integrity Preliminary Report](https://gitlab.com/ccao-data-science---modeling/ccao_sf_cama_dev/-/blob/master/documentation/Preliminary%20Report%20on%20Data%20Integrity%20June%207,%202019.pdf). In Project A2, you will be building a linear regression model that predicts sales prices using training data, but it's important to first understand how the structure of the data informs such a model. In this section, we will create a sequence of exploratory visualizations and perform feature engineering in preparation for that prediction task.

Note that we will perform EDA on the **initial data**.

### 🔖 Sale Price
We begin by examining the distribution of our target variable `Sale Price`. We have provided the following helper method `plot_distribution` that you can use to visualize the distribution of the `Sale Price` using both the histogram and the box plot using a shared x-axis. Run the following 2 cells.

```python
def plot_distribution(data, label):
    fig, axs = plt.subplots(nrows=2)

    sns.distplot(
        data[label], 
        ax=axs[0]
    )
    sns.boxplot(
        x=data[label],
        width=0.3, 
        ax=axs[1],
        showfliers=False,
    )

    # Align axes
    spacer = np.max(data[label]) * 0.05
    xmin = np.min(data[label]) - spacer
    xmax = np.max(data[label]) + spacer
    axs[0].set_xlim((xmin, xmax))
    axs[1].set_xlim((xmin, xmax))

    # Remove some axis text
    axs[0].xaxis.set_visible(False)
    axs[0].yaxis.set_visible(False)
    axs[1].yaxis.set_visible(False)

    # Put the two plots together
    plt.subplots_adjust(hspace=0)
    fig.suptitle("Distribution of " + label)
```

```python
plot_distribution(initial_data, label='Sale Price')
```

We can see that the data is highly skewed to the right currently, probably because of some extremely expensive houses. To get a clearer view of the distribution, run the following cell to examine a more limited range of house prices from [\\$0, \\$1,000,000]. 

```python
no_right_outliers = initial_data[initial_data['Sale Price'] <= 1000000]['Sale Price']
plt.hist(no_right_outliers, bins=200)
plt.xlabel("Sale Price")
plt.ylabel("Count")
plt.title("Distribution of Sale Price")
plt.show()
```

Interestingly, we also see this spike in the smallest bin of our `Sale Price` distribution, something that wasn't immediately evident from our initial plot. Run the following cell to see which values are the most common.

```python
initial_data['Sale Price'].value_counts()
```

This should be immediately concerning. While we don't know the exact reason for so many houses being reported as \\$1 (most likely a placeholder value of some sort), including these data points in any model would heavily skew the fit. After all, most houses are worth more than \\$1!

<br><br>

---
### 🚀 Question 2
In this question, you will explore small values of `Sale Price` in our data.

<!-- BEGIN QUESTION -->

#### Question 2a
In the cell below, write an analysis plan to identify the four most frequent `Sale Price` values in the range [\\$0, \\$500). Then, use an LLM to implement your plan and assign the results to `q2a` as a list of integers ordered from most to least frequent.


```python
# Write your analysis plan here.
...
q2a = ...
```

```python
grader.check("q2a")
```

<!-- END QUESTION -->

This exploration has indicated to us that, moving forward, we should probably pick a cutoff to get rid of smaller, likely unrepresentative `Sale Price`s. As a result, moving forward you will be asked to only look at households whose price is at least \\$500 dollars. While this cutoff is slightly arbitrary, it is sufficient for our purpose of eliminating unrealistically low values.

On the other hand, there are also the upper-bound outliers that skewed our data in the first place. Run the cell below to find the proportion of buildings sold for over \\$1,000,000.

```python
len(initial_data[initial_data['Sale Price'] > 1000000]) / len(initial_data[initial_data['Sale Price'] >= 500])
```

Only about 3% of properties—roughly 5,000—sold for more than \\$1,000,000, so restricting the model to sale prices between \\$500 and \\$1,000,000 would be reasonable. In Project A1, however, we will model the full range of sale prices ([\\$500, infinity)). As a result, we will need to reason about how these outliers shape the distribution of points. One way we can make our initial visualization more legible is by log-transforming our data to minimize the effect of these large outliers.

<!-- BEGIN QUESTION -->

#### Question 2b
In the previous part, we saw that many rows have a sale price of \\$1. Using external sources and a bit of research, explain why these prices are so low, and what's actually happening with these houses. Based on your findings, explain why it's a **good idea** to remove these rows from our dataset.

- **Hint**: *This is not a marketing trick to make the house seem more appealing!* 

_Type your answer here, replacing this text._

```python
# CELL FOR SCRATCH WORK / EDA
```

<!-- END QUESTION -->

<br><br>

As previously established to zoom in on the visualization of most households, we will focus only on a subset of `Sale Price` for this assignment. In addition, it may be a good idea to apply a log transformation to `Sale Price`. 

We will create the `training_data` DataFrame by modifying the `initial_data` DataFrame.
- `training_data` will contain only households whose price is at least $500.
- `training_data` will contain a new `Log Sale Price` column that contains the log-transformed sale prices.

**Note**: This also implies that, from now on, our target variable in the model will be the log-transformed sale prices from the column `Log Sale Price`.

```python
training_data = initial_data[initial_data['Sale Price'] >= 500]
training_data['Log Sale Price'] = np.log(training_data['Sale Price'])
```

Let's create a new distribution plot using the log-transformed sale prices. As a sanity check, you should see that the distribution for `Log Sale Price` is much more uniform.

```python
plot_distribution(training_data, label='Log Sale Price');
```

<br><br>

---

Next, we want to explore whether there is any correlation between `Log Sale Price` and the total area occupied by the property. The `codebook.txt` file tells us the column `Building Square Feet` should do the trick — it measures "(from the exterior) the total area, in square feet, occupied by the building".

Similar to `Sale Price`, `Building Square Feet` may also benefit from a log transformation. In the following sections, we will explore `Log Building Square Feet`.

```python
training_data['Log Building Square Feet'] = np.log(training_data['Building Square Feet'])
```

<!-- BEGIN QUESTION -->

<br><br>

---
### 🚀 Question 3a

In the visualization below, we created a `jointplot` with `Log Building Square Feet` on the x-axis, and `Log Sale Price` on the y-axis. In addition, we fit a simple linear regression line through the bivariate scatter plot in the middle.

Based on the following plot, would `Log Building Square Feet` make a good candidate as one of the features for our model? Why or why not? 

**Be specific**: Your answer should include "If the graph looked like \<specific description of change here\>, then I would answer the opposite way."

**Hint:** To help answer this question, ask yourself: What kind of relationship does a “good” feature share with the target variable we aim to predict?

![Joint Plot](images/q2p3_jointplot.png)


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br><br>

---
### 🚀 Question 3b
In the visualization below, we have plotted a heatmap displaying the correlations between various features, one of which is `Log Sale Price`.

Based on this heatmap, rank the provided features by how informative you think they'll be in predicting `Log Sale Price` from most informative to least informative. Justify your answer.

- **Hint**: remember what you know about correlation, and notice that we used `center=0` when creating the heatmap so that values close to 0 will be at the center of the color scale.

```python
q3b_features = ['Log Sale Price', 'Building Square Feet', 'Most Recent Sale', 'Age', 'Road Proximity', 'Sale Month of Year']
sns.heatmap(data=training_data[q3b_features].corr(numeric_only=True), annot=True, center=0)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

---

Continuing from the previous part, as you explore the dataset, you might still run into more outliers that prevent you from creating a clear visualization or capturing the trend of the majority of the houses. 

Below, we defined a function `remove_outliers` that removes outliers from the dataset based on a threshold value of a variable. For example, `remove_outliers(training_data, 'Building Square Feet', lower=500, upper=8000)` should return a copy of `data` with only observations that satisfy `Building Square Feet` less than 8000 (exclusive) and `Building Square Feet` greater than 500 (exclusive).

<!-- grade_id: cell-9186ec2ca053d0aa, points: 0, grade: false, solution: true -->
```python
def remove_outliers(data, variable, lower=-np.inf, upper=np.inf):
    """
    Input:
      data (DataFrame): the table to be filtered
      variable (string): the column with numerical outliers
      lower (numeric): observations with values lower than or equal to this will be removed
      upper (numeric): observations with values higher than or equal to this will be removed
    
    Output:
      a DataFrame with outliers removed
      
    Note: This function should not change or mutate the contents of data.
    """  

    return data[(data[variable] < upper) & (data[variable] > lower)]
```

<!-- BEGIN QUESTION -->

<br><br>

---
### 🚀 Question 4a

To give an example of how `remove_outliers` might be useful in the future, plot the distribution of the attribute `Estimate (Land)` below using `plot_distribution` with the provided code below. From `codebook.txt`, `Estimate (Land)` estimates the market value of a property’s land from the prior tax year. 

Indicate whether the following statement is correct. Assign your answer to `q4astatement`. 

“At least 95% of the properties in `training_data` have `Estimate (Land)` values less than $500,000.00”

```python
# Run this to see the distribution!
plot_distribution(training_data, label='Estimate (Land)');
```

```python
# This should be set to True or False
q4astatement = ...
```

```python
grader.check("q4a")
```

<!-- END QUESTION -->

<br><br>

---

Similar to the exploration of `Sale Price` in Question 2, you may have observed that `Estimate (Land)` contains some points that heavily skew its distribution. One way to address these issues is to transform the data to a different scale, as we did with `Sale Price`. Another way is to remove outliers. 

Here, we define an outlier as any point that is 1.5 interquartile ranges above Q3 or below Q1 (exclusive of the boundary).

If you need a refresher on how the IQR works, you can read this [Wikipedia page on IQR](https://en.wikipedia.org/wiki/Interquartile_range).

```python
IQR = training_data['Estimate (Land)'].quantile(0.75) - training_data['Estimate (Land)'].quantile(0.25) 

upper_boundary = training_data['Estimate (Land)'].quantile(0.75) + (IQR * 1.5) 
lower_boundary = training_data['Estimate (Land)'].quantile(0.25) - (IQR * 1.5) 

filtered_training_data = remove_outliers(training_data, "Estimate (Land)", lower = lower_boundary, upper = upper_boundary)
```

Let's create a new distribution plot using the `Estimate (Land)` with outliers removed. You should see that the distribution for `Estimate (Land)` is much less skewed with `filtered_training_data`!

```python
plot_distribution(filtered_training_data, label='Estimate (Land)');
```

<!-- BEGIN QUESTION -->

<br><br>

---
### 🚀 Question 4b
Earlier in the assignment, we removed outliers to help us visualize the house price and land area values for the vast majority of homes. But, removing outliers from our training data could create problems downstream. Give two **hypothetical** examples of houses that would get worse predictions from your model if outliers are removed. Be specific: describe these houses in as much detail as necessary to explain why they would be affected.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/>
<hr style="border: 1px solid #fdb515;" />

## ⛲ Part 3: Feature Engineering

In this section, we will walk you through a few feature engineering techniques. 

### 🚽 Bathrooms

Let's start simple by extracting the total number of bathrooms as our first feature for the model. You may notice that the `Bathrooms` column doesn't actually exist in the original `DataFrame`! Instead, it is part of the `Description` column.

<br><br>

---
### 🚀 Question 5a

Let's take a closer look at the `Description` column first. Compare the description for a few rows. For the following list of variables, how many of them can be extracted from the `Description` column? Assign your answer to a list of integers corresponding to the statements that you think are true (ie. `[1, 2, 3]`).

1. The date the property was sold on.
2. The previous owner of the property.
3. The number of garages the property has.
4. The total number of bedrooms inside the property.
5. The number of stories the property contains.
6. The address of the property.
7. The total number of rooms inside the property.
8. The total number of bathrooms inside the property.

```python
# optional cell for scratch work 
```

```python
q5a = ...
```

```python
grader.check("q5a")
```

```python
# optional cell for scratch work
```

<br><br>

---
### 🚀 Question 5b

Write a `regex` pattern that extracts the number of bathrooms for each house, and assign your answer to `rooms_regex`.

**Hint**: You should consider inspecting the `Description` column to figure out whether there is any general structure within the text. Once you have noticed a certain pattern, you are set with the power of `regex`!

```python
def add_total_bathrooms(data):
    """
    Input:
      data (DataFrame): a DataFrame containing at least the Description column.

    Output:
      a DataFrame with a new column "Bathrooms" containing floats.

    """
    with_rooms = data.copy()
    ...
    rooms = with_rooms['Description'].str.extract(rooms_regex).astype(float)
    with_rooms['Bathrooms'] = rooms
    
    return with_rooms

training_data = add_total_bathrooms(training_data)
```

```python
grader.check("q5b")
```

<!-- BEGIN QUESTION -->

<br><br>

---
### 🚀 Question 5c

Create a visualization that clearly and succinctly shows whether there exists an association between  `Bathrooms` and `Log Sale Price`. A good visualization should satisfy the following requirements:
- It should avoid overplotting.
- It should have clearly labeled axes and a succinct title.
- It should convey the strength of the correlation between `Log Sale Price` and the number of `Bathrooms`: in other words, you should be able to look at the plot and describe the general relationship between `Log Sale Price` and `Bathrooms`.

Start by describing your visualization in words. Then, write an analysis plan to create the data frame needed for your visualization. Finally, use an LLM to translate your analysis plan to code.

We will only grade the code that produces your visualization.

**Hint 1**: A direct scatter plot of the `Log Sale Price` against the number of rooms for all of the households in our training data might risk overplotting. Try it out and see why it causes overplotting.

**Hint 2**: Take a closer look at the variable `Bathrooms`. What variable type is `Bathrooms` and what are the possible values `Bathrooms` can take? What visualizations might be good at showing this?

_Type your answer here, replacing this text._

```python
...
```

<!-- END QUESTION -->

Now, let's take a look at the relationship between neighborhood and sale prices of the houses in our dataset.
Notice that currently we don't have the actual names for the neighborhoods. Instead we will use a similar column, `Neighborhood Code` (which is a numerical encoding of the actual neighborhoods by the Assessor’s office).

<br><br>

---
### 🚀 Question 6a

If we try directly plotting the distribution of `Log Sale Price` for all of the households in each neighborhood using the `plot_categorical` function from the next cell, we get the following visualization.


![overplot](images/q5p2_catplot.png)


```python
# Feel free to create a cell below this and run plot_categorical(training_data) if you want to see what this function outputs.
def plot_categorical(neighborhoods):
    fig, axs = plt.subplots(nrows=2)

    sns.boxplot(
        x='Neighborhood Code',
        y='Log Sale Price',
        data=neighborhoods,
        ax=axs[0],
    )

    sns.countplot(
        x='Neighborhood Code',
        data=neighborhoods,
        ax=axs[1],
    )

    # Draw median price
    axs[0].axhline(
        y=training_data['Log Sale Price'].median(), 
        color='red',
        linestyle='dotted'
    )

    # Label the bars with counts
    for patch in axs[1].patches:
        x = patch.get_bbox().get_points()[:, 0]
        y = patch.get_bbox().get_points()[1, 1]
        axs[1].annotate(f'{int(y)}', (x.mean(), y), ha='center', va='bottom')

    # Format x-axes
    axs[1].set_xticklabels(axs[1].xaxis.get_majorticklabels(), rotation=90)
    axs[0].xaxis.set_visible(False)

    # Narrow the gap between the plots
    plt.subplots_adjust(hspace=0.01)
```

<!-- BEGIN QUESTION -->

Oh no, looks like we have run into the problem of overplotting again! 

You might have noticed that the graph is overplotted because **there are actually quite a few neighborhoods in our dataset**! For the clarity of our visualization, we will have to zoom in again on a few of them.

Write an analysis plan to identify the 10 `Neighborhood Code` values with the most property sales. Then, use an LLM to implement your plan and assign a DataFrame containing only properties in those neighborhoods to `in_top_10_neighborhoods`.

```python
# Write your analysis plan here.
top_10_neighborhood_codes = ...
in_top_10_neighborhoods = ...
```

```python
grader.check("q6a")
```

<!-- END QUESTION -->

Let's create another visualization of sale price within in each neighborhood again, but this time with a narrower focus!

```python
plot_categorical(neighborhoods=in_top_10_neighborhoods)
```

<!-- BEGIN QUESTION -->

<br><br>

---
### 🚀 Question 6b

The plot above shows that the number of data points varies substantially across neighborhoods. For example, Neighborhood 200 has only around 44% of the number of datapoints as Neighborhood 30.

One way to address differences in the amount of data available across neighborhoods is to create a feature that identifies properties in higher-priced neighborhoods. In this question, you will determine which `n` neighborhoods should be considered the most expensive. In `q6c`, you will use the result to label properties located in those neighborhoods.

Write an analysis plan that starts with property-level data and ends with a list of the top `n` most expensive neighborhood codes, based on `metric`. Then, use an LLM to implement your plan in `find_expensive_neighborhoods`.

For example, calling `find_expensive_neighborhoods(training_data, n=3, metric=np.median)` should return the 3 neighborhood codes with the highest median `Log Sale Price` computed across all properties in those neighborhood codes. 

```python
# Write your analysis plan here.
def find_expensive_neighborhoods(data, n=3, metric=np.median):
    """
    Input:
      data (DataFrame): should contain at least an int-valued 'Neighborhood Code'
        and a numeric 'Log Sale Price' column
      n (int): the number of top values desired
      metric (function): function used for aggregating the data in each neighborhood.
        for example, np.median for median prices
    
    Output:
      a list of the neighborhood codes of the top n highest-priced neighborhoods 
      as measured by the metric function
    """
    neighborhoods = ...
    
    # This makes sure the final list contains the generic int type used in Python3, not specific ones used in NumPy.
    return [int(code) for code in neighborhoods]

expensive_neighborhoods = find_expensive_neighborhoods(training_data, 3, np.median)
expensive_neighborhoods
```

```python
grader.check("q6b")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br><br>

---
### 🚀 Question 6c

Write an analysis plan to add an integer column named `in_expensive_neighborhood` to `data`. The column should contain 1 if a property’s neighborhood is in `expensive_neighborhoods` and 0 otherwise. Then, use an LLM to implement your plan in `add_in_expensive_neighborhood`.

```python
# Write your analysis plan here.
def add_in_expensive_neighborhood(data, expensive_neighborhoods):
    """
    Input:
      data (DataFrame): a DataFrame containing a 'Neighborhood Code' column with values
        found in the codebook
      expensive_neighborhoods (list of ints): ints should be the neighborhood codes of
        neighborhoods pre-identified as expensive
    Output:
      DataFrame identical to the input with the addition of a binary
      in_expensive_neighborhood column
    """
    return data

expensive_neighborhoods = find_expensive_neighborhoods(training_data, 3, np.median)
training_data = add_in_expensive_neighborhood(training_data, expensive_neighborhoods)
```

```python
grader.check("q6c")
```

<!-- END QUESTION -->

In the following section, we will take a closer look at the `Wall Material` feature of the dataset and examine how we can incorporate categorical features into our linear model.

<br><br>

---

If we look at `codebook.txt` carefully, we can see that the Assessor's Office uses the following mapping for the numerical values in the `Wall Material` column.
```
Wall Material: 

       1    Wood
       2    Masonry
       3    Wood&Masonry
       4    Stucco

```

We wrote the function `substitute_wall_material` that replaces each numerical value in `Wall Material` with their corresponding wall material.

```python
def substitute_wall_material(data):
    """
    Input:
      data (DataFrame): a DataFrame containing a 'Wall Material' column.  Its values
                         should be limited to those found in the codebook
    Output:
      new DataFrame identical to the input except with a refactored 'Wall Material' column
    """
    replacements = {
        'Wall Material': {
            1: 'Wood',
            2: 'Masonry',
            3: 'Wood&Masonry',
            4: 'Stucco'
        }
    }
    
    new_data = data.replace(replacements)
    return new_data
    
training_data_mapped = substitute_wall_material(training_data)
training_data_mapped.head()
```

<br><br>

---

#### An Important Note on One-Hot-Encoding 

Unfortunately, simply replacing the integers with the appropriate strings isn’t sufficient for using `Wall Material` in our model.  Since `Wall Material` is a categorical variable, we will have to one-hot-encode the data. [For more information on why we want to use one-hot-encoding, refer to this link](https://machinelearningmastery.com/why-one-hot-encode-data-in-machine-learning/).

We have provided a prebuilt ohe_wall_material function below. It takes a DataFrame containing the labeled `Wall Material` column and returns a new DataFrame with all original columns plus one binary indicator column for each wall material category. The new columns follow the naming format Wall `Material_MATERIAL`. Run the cell below to define and apply the function.

```python
from sklearn.preprocessing import OneHotEncoder

def ohe_wall_material(data):
    """
    One-hot-encodes wall material. New columns are of the form "Wall Material_MATERIAL".
    """
    
    oh_enc = OneHotEncoder()
    oh_enc.fit(data[['Wall Material']])
    oh_enc_columns = pd.DataFrame(oh_enc.transform(data[['Wall Material']]).toarray(), # .todense() also works
                           columns=oh_enc.get_feature_names_out(),
                           index = data.index)
    
    return data.merge(oh_enc_columns, left_index=True, right_index=True)
    

training_data_ohe = ohe_wall_material(training_data_mapped)
# This line of code will display only the one-hot-encoded columns in training_data_ohe that 
# have names that begin with “Wall Material_" 
training_data_ohe.filter(regex='^Wall Material_').head(10)
```

<!-- BEGIN QUESTION -->

<br><br>

---
### 🚀 Question 7a
So far, we've been exploring different features to pass into our model, but we have yet to explore the effect of different error metrics on our model. Let's get some practice interpreting different error metrics.

**Your task**: Describe/write in mathematical notation a loss function that assigns higher penalty to errors on expensive houses. You are free to choose exactly how it does so. Additionally, calculate the loss function's value for two cases: 

1. A house that sold for \\$2M where the model predicted \\$2.2M.
2. A house that sold for \\$750K where the model predicted \\$950K.

Explain why your loss function produces a result that makes sense, and if it doesn't, consider modifying your answer so that it does.

- Please write your loss function in LaTeX. You can do this inline by surrounding your latex expression by two dollar sign characters \\$. For example, (x - y)^2 will appear as $ (x - y)^2 $ when written between two dollar signs.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br><br>

---
### 🚀 Question 7b
In this part, you'll interpret a loss function. Consider the expression for the loss function below, where $ \alpha > \beta > 0 $:

$$
\text{FloorCeilPen} = \frac{1}{n} \sum_{i=1}^n
\begin{cases}
\alpha \cdot |y_i - \hat{y}_i| & \text{if } \hat{y}_i < y_i \\
\beta \cdot |y_i - \hat{y}_i| & \text{if } \hat{y}_i \geq y_i
\end{cases}
$$

**Your task**: Describe what this loss function does in plain English.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

---
## Final Note
Note that for Project A1, we have released guides for the hidden tests on the [Debugging Guide](https://ds100.org/debugging-guide/proja1-hidden-test/). Please set `hidden_tests_ack` to `True` below once you've read this!

```python
hidden_tests_ack = ...
```

```python
grader.check("final_note")
```

<br>

---

## Takeaways

In this project, you have...
- Performed basic Exploratory Data Analysis (EDA)
- Added new features to the data
- Understood the effects of different error metrics

### What's next? 

In Project A2, you will focus on building a linear regression model to predict house prices. You are well-prepared to build such a model: You have considered what is in this dataset, what it can be used for, and engineered some features that should be useful for prediction. As we look into creating a house-pricing model for Cook County, we'll also consider some challenging social implications this model may carry. 

## Course Content Feedback

If you have any feedback about this assignment or about any of our other weekly assignments, lectures, or discussions, please fill out the [Course Content Feedback Form](https://docs.google.com/forms/d/e/1FAIpQLSeNCHbTUipouL3yWFeRdl9Blwu-U4zSciTi4Kcb1p5ykCtb-A/viewform). Your input is valuable in helping us improve the quality and relevance of our content to better meet your needs and expectations!

## Submission Instructions

Below, you will see a cell. Running this cell will automatically generate a zip file with your autograded answers. Once you submit this file to the Project A1 Coding assignment on Gradescope, Gradescope will automatically submit a PDF file with your written answers to the Project A1 Written assignment. If you run into any issues when running this cell, feel free to check this [section](https://ds100.org/debugging-guide/autograder-gradescope/#why-does-the-last-grader-export-cell-fail-if-all-previous-tests-passed) in the Data 100 Debugging Guide.

If there are issues with automatically generating the PDF, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit those.

Additionally, make sure you have submitted your answers to `q1bii` and `q1cii` as a pdf to the Project A1 Plot Sketch assignment on Gradescope.

**You are responsible for ensuring your submission follows our requirements, that the PDF for Proj A1 written answers was generated/submitted correctly, and that the file was submitted to the correct assignment. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline. 

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for final_note (points: 0)
```
>>> hidden_tests_ack == True
True
```


Tests for q1bii (points: 0)
```
>>> q1bii == True
True
```


Tests for q1cii (points: 0)
```
>>> q1cii == True
True
```


Tests for q2a (points: 1)
```
>>> isinstance(q2a, list) and len(q2a) >= 3 and isinstance(q2a[0], int)
True
```


Tests for q4a (points: 1)
```
>>> q4astatement in [True, False]
True
```


Tests for q5a (points: 1)
```
>>> bool(np.all([num in np.arange(1, 9) for num in q5a]))
True
```
```
>>> isinstance(q5a, list) == True
True
```
```
>>> isinstance(q5a[0], int) == True
True
```


Tests for q5b (points: 2)
```
>>> bool(not training_data['Bathrooms'].isnull().any())
True
```
```
>>> bool(pd.api.types.is_float_dtype(training_data['Bathrooms']))
True
```


Tests for q6a (points: 2)
```
>>> bool(len(in_top_10_neighborhoods['Neighborhood Code'].unique()) <= 20)
True
```
```
>>> bool(in_top_10_neighborhoods['Neighborhood Code'].iloc[0] == 120)
True
```


Tests for q6b (points: 2)
```
>>> bool(len(find_expensive_neighborhoods(training_data, 5, np.median)) == 5)
True
```
```
>>> bool(isinstance(expensive_neighborhoods, list))
True
```
```
>>> bool(set(find_expensive_neighborhoods(training_data, 2, np.min)) == set([106, 580]))
True
```


Tests for q6c (points: 1)
```
>>> bool(sum(training_data.loc[:, 'in_expensive_neighborhood']) == 1290 and sum(training_data.loc[:, 'in_expensive_neighborhood'].isnull()) == 0)
True
```
